```python
# Initialize Otter
import otter
grader = otter.Notebook("quiz8practice-functions-loops.ipynb")
```

# ⌛️ Quiz 8 Practice - Using Functions

This quiz will evaluate your mastery of using functions in Python. Functions provide a way to isolate code that you want to use repeatedly, and they allow you to pass in data to the code and get data back out of the code.

**Note** The actual quiz will have different questions but will be similar to the ones provided here. All tests are visible to you, so you can test your solution as many times as you want.


## Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
import pkg_resources
from subprocess import call

package_name = 'ENGR131_Util_2024'
version = '0.1.8'
package_version = f'{package_name}=={version}'

try:
    # Check if the package and version are installed
    pkg_resources.require(package_version)
    print(f'{package_version} is already installed.')
except pkg_resources.DistributionNotFound:
    # If not installed, install the package
    print(f'{package_version} not found. Installing...')
    call(['pip', 'install', package_version])
except pkg_resources.VersionConflict:
    # If a different version is installed, you can choose to upgrade/downgrade
    installed_packages = {dist.key: dist.version for dist in pkg_resources.working_set}
    installed_version = installed_packages.get(package_name.lower())
    print(f'{package_name} {installed_version} is installed, but {version} is required.')
    # Optionally, upgrade or downgrade the package to the required version
    call(['pip', 'install', '--upgrade', package_version])

```

```python
# DO NOT MODIFY THIS CELL

from ENGR131_Util_2024 import cell_logger, StudentInfoForm, responses, upsert_to_json_file
# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', cell_logger)
responses["assignment"] = "quiz_8_practice"

StudentInfoForm(**responses)
```

**Question 1: Thermal Expansion of a Rod**

You are a materials engineer tasked with evaluating the suitability of a metal rod for use in a construction project where temperature variations are significant. It's essential to estimate the thermal expansion of the rod to ensure that it will not cause structural issues due to temperature changes.

The change in length, $\Delta L$, of a rod due to thermal expansion can be calculated using the equation:

$$
\Delta L = \alpha L_0 \Delta T \tag{1}
$$

where:
- $\alpha$ is the coefficient of thermal expansion of the material (in per degree Celsius, $°C^{-1}$),
- $L_0$ is the original length of the rod (in meters, m),
- $\Delta T$ is the change in temperature (in degrees Celsius, °C).

To ensure the structural integrity of the construction, it's crucial to consider a tolerance factor in the design. This factor adjusts the expected change in length to account for uncertainties in material properties and temperature variations.

If the change in length is $\Delta L$, the adjusted length change including the tolerance factor is given by:

$$ \Delta L_{\text{adjusted}} = \text{Tolerance Factor} \times \Delta L \tag{2}$$

Write Python code to do the following:

* Define a function called `rod_expansion` which accepts three input arguments: `alpha`, `L0`, and `DeltaT`, and a fourth optional argument called `tolerance_factor` which defaults to 1.0.
* Define a variable called `DeltaL` and set it equal to the result of equation (1) above.
* Define a variable called `DeltaL_adjusted` and set it equal to the result of the equation (2) above.
* Have your function print the response: `The change in length of the rod is: <DeltaL> m, and the adjusted length change for tolerance is <DeltaL_adjusted> m`, where `<DeltaL>`, and `<DeltaL_adjusted>` are the float values of the variable for the change in length and the adjusted change respectively, rounded to 3 decimal places.
* Implement your function so that it returns the adjusted length change of the rod for a given input.

Provide an example usage of your function by calling it with the following inputs:
* `alpha = 1.2e-5` °C\(^{-1}\)
* `L0 = 2.0` m
* `DeltaT = 50` °C
* `tolerance_factor = 10`

Save the output of the function call to a variable `adjusted_length_change`.

Your code should print the response: 

Note: You can control the precision of the output by using string formatting to round the output to a specific number of decimal places. For a consistent presentation, use three decimal places in your output as shown in the example:

```python
print(f"{DeltaL:0.3f}")
```

Your code replaces the prompt:  `...`

```python
# Your function for rod_expansion goes here
...

# Example usage:
alpha = ...
L0 = ...
DeltaT = ...
tolerance_factor = ...

# Call the function with example inputs and save to a variable `adjusted_length_change`
...
```

```python
grader.check("q1-thermal-expansion-rod")
```

**Question 2: Identifying the First Acidic Solution with pH Below a Certain Level 5**

Write a Python function `acidic_pH_detection` that takes a sequence of pH level measurements (represented as a list of floats) as a variable `pH_levels`. The goal is to identify the first measurement that falls below a predefined `threshold` value, indicating the occurrence of an acidic solution. Once an acidic solution with a pH below the threshold is found, the program should return the *index* (position) of this measurement in the dataset and the *value* of the pH level, as a tuple then terminate the search.

Requirements:
1. Write a function `acidic_pH_detection` that takes two input arguments: `pH_levels` (a list of floats) and `threshold` (a float).

   1. Use a for loop to iterate through the `pH_levels` list.
   2. Within the loop, implement a condition to check if the current pH level is below the threshold.
   3. Print the index of the first pH level that is below the threshold and the value of this pH level. The print statement should read `First acidic solution detected below threshold found at index <index> with pH <value>.`, where <index> and <value> are the index and pH level of the first measurement below the threshold.
   4. If a pH level below the threshold is found, use a return statement to exit the loop immediately and return a tuple containing the index of the measurement and the value of the pH level.
   5. If no pH level below the threshold is detected, print a message: `No acidic solution was detected below the threshold of <threshold>.`, where <threshold> is the value of the threshold.

2. Test your function by calling it with the provided inputs. Save the output of the function call to a variable `acidic_solution`.

Hints:
* Use the `enumerate()` function to get both the index and the value of each item when iterating through the list with a for loop.

```python
# Sample pH level measurements
pH_levels = [7.2, 5.5, 6.8, 7.0, 7.3, 4.8, 3.2, 7.4, 8.0, 6.5, 4.9, 3.5]

# Threshold value for acidity
threshold_pH = 5.0

...
```

```python
grader.check("q2-detecting-threshold-signal")
```

## Fin

Tests for q1-thermal-expansion-rod (points: 12)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_1')
>>> drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_2')
>>> drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_3')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q1_4')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q1_5')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q1_6')
>>> drexel_jupyter_logger.variable_logger_csv('0, 3', 'q1_7')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q2_1')
>>> drexel_jupyter_logger.variable_logger_csv('0, 2', 'q2_2')
>>> drexel_jupyter_logger.variable_logger_csv('0, 3', 'q2_3')
>>> drexel_jupyter_logger.variable_logger_csv('0, 3', 'q2_4')
>>> drexel_jupyter_logger.variable_logger_csv('0, 4', 'q2_5')
>>> import json
>>> student_data = {'first_name': responses['first_name'], 'last_name': responses['last_name'], 'drexel_id': responses['drexel_id'], 'drexel_email': responses['drexel_email']}
>>> with open('student_data.json', 'w') as json_file:
...     json.dump(student_data, json_file)
>>> scorer = submit_score()
>>> question_id = 'q1_1'
>>> if alpha == 1.2e-05 and L0 == 2.0 and (DeltaT == 50) and (tolerance_factor == 10):
...     drexel_jupyter_logger.variable_logger_csv('2, 1', question_id)
...     score = 1
... else:
...     score = 0
>>> with patch('builtins.print') as mock_print:
...     response = {'question_id': question_id, 'score': score, 'max_score': 1, 'student_response': f'{alpha}, {L0}, {DeltaT}, {tolerance_factor}'}
...     scorer.add_response(response)
...     scorer.submit()
>>> assert alpha == 1.2e-05, 'The value of alpha is not correct.'
>>> assert L0 == 2.0, 'The value of L0 is not correct.'
>>> assert DeltaT == 50, 'The value of DeltaT is not correct.'
>>> assert tolerance_factor == 10, 'The value of the tolerance factor is not correct.'
```
success_message: Test values are implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> import numpy as np
>>> from unittest.mock import patch
>>> expected_adjusted_length_change = 0.012
>>> scorer = submit_score()
>>> question_id = 'q1_2'
>>> score = 1 if np.isclose(adjusted_length_change, expected_adjusted_length_change) else 0
>>> with patch('builtins.print') as mock_print:
...     response = {'question_id': question_id, 'score': score, 'max_score': 1, 'student_response': f'{adjusted_length_change}'}
...     scorer.add_response(response)
...     scorer.submit()
>>> assert np.isclose(adjusted_length_change, expected_adjusted_length_change), 'The adjusted length change is not calculated correctly.'
```
success_message: Test for adjusted length change implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_3'
>>> with patch('builtins.print') as mock_print:
...     assert type(rod_expansion(alpha=1.2e-05, L0=1.0, DeltaT=30, tolerance_factor=1.0)) == float, 'The function does not return a numerical value of type float.'
...     score = 1 if type(rod_expansion(alpha=1.2e-05, L0=1.0, DeltaT=30, tolerance_factor=1.0)) == float else 0
...     response = {'question_id': question_id, 'score': score, 'max_score': 1, 'student_response': f'{type(rod_expansion(alpha=1.2e-05, L0=1.0, DeltaT=30, tolerance_factor=1.0))}'}
...     scorer.add_response(response)
...     scorer.submit()
```
success_message: Function correctly returns a numerical value.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_4'
>>> (alpha, L0, DeltaT, tolerance_factor) = (2000, 0.4, 0.002, 4)
>>> with patch('builtins.print') as mock_print:
...     assert rod_expansion(alpha, L0, DeltaT) == 1.6, 'thermal expansion with default tolerence factor is not correct.'
...     if rod_expansion(alpha, L0, DeltaT) == 1.6:
...         score = 2
...         drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
...     else:
...         score = 0
...     response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{rod_expansion(alpha, L0, DeltaT)}'}
...     scorer.add_response(response)
...     scorer.submit()
```
success_message: default tolerance factor is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_5'
>>> (alpha, L0, DeltaT, tolerance_factor) = (2000, 0.4, 0.002, 4)
>>> with patch('builtins.print') as mock_print:
...     out = rod_expansion(alpha, L0, DeltaT, tolerance_factor)
...     assert out == 6.4, 'tolerance factor implemented correctly.'
...     if out == 6.4:
...         score = 2
...         drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
...     else:
...         score = 0
...     response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{out}'}
...     scorer.add_response(response)
...     scorer.submit()
```
success_message: tolerance factor is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> import inspect
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_6'
>>> signature = inspect.signature(rod_expansion)
>>> parameters = signature.parameters
>>> assert len(parameters) == 4, f'Expected 4 parameters, found {len(parameters)}'
>>> tolerance_factor_param = parameters.get('tolerance_factor')
>>> assert tolerance_factor_param is not None, 'tolerance factor parameter missing.'
>>> assert tolerance_factor_param.default != inspect.Parameter.empty, 'tolerance factor does not have a default value, thus not optional.'
>>> if len(parameters) == 4 and tolerance_factor_param is not None and (tolerance_factor_param.default != inspect.Parameter.empty):
...     score = 2
...     drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
... else:
...     score = 0
>>> with patch('builtins.print') as mock_print:
...     response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{len(parameters)}, {tolerance_factor_param.default}'}
...     scorer.add_response(response)
...     scorer.submit()
```
success_message: Input parameters is is correctly defined.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_7'
>>> expected_message = 'The change in length of the rod is: 0.001 m, and the adjusted length change for tolerance is 0.012 m'
>>> with patch('builtins.print') as mock_print:
...     alpha = 1.2e-05
...     L0 = 2.0
...     DeltaT = 50
...     tolerance_factor = 10
...     out = rod_expansion(alpha, L0, DeltaT, tolerance_factor)
...     mock_print.assert_called_once_with(expected_message)
>>> if mock_print.call_args[0][0] == expected_message:
...     score = 3
...     drexel_jupyter_logger.variable_logger_csv('3, 3', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 3, 'student_response': f'{mock_print.call_args[0][0]}'}
>>> with patch('builtins.print') as mock_print:
...     scorer.add_response(response)
...     scorer.submit()
```
success_message: Input parameters is is correctly defined.


Tests for q2-detecting-threshold-signal (points: 14)
```
>>> import drexel_jupyter_logger
>>> import inspect
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q2_1'
>>> signature = inspect.signature(acidic_pH_detection)
>>> parameters = signature.parameters
>>> if len(parameters) == 2 and 'acidic_pH_detection' in locals() and callable(acidic_pH_detection):
...     score = 2
...     drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{len(parameters)}, {callable(acidic_pH_detection)}'}
>>> with patch('builtins.print') as mock_print:
...     scorer.add_response(response)
...     scorer.submit()
>>> assert len(parameters) == 2, f'Expected 2 parameters, found {len(parameters)}'
>>> assert 'acidic_pH_detection' in locals(), 'Function acidic_pH_detection is not defined.'
>>> assert callable(acidic_pH_detection), 'acidic_pH_detection is not callable.'
```
success_message: Function correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_2'
>>> expected_message = 'First acidic solution detected below threshold found at index 3 with pH 4.8.'
>>> with patch('builtins.print') as mock_print:
...     new_pH_levels = [6.5, 6.7, 6.8, 4.8, 5.1, 5.5]
...     (index, pH) = acidic_pH_detection(new_pH_levels, threshold_pH)
...     mock_print.assert_called_once_with(expected_message)
>>> if mock_print.call_args[0][0] == expected_message:
...     score = 2
...     drexel_jupyter_logger.variable_logger_csv('2, 2', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 2, 'student_response': f'{mock_print.call_args[0][0]}'}
>>> with patch('builtins.print') as mock_print:
...     scorer.add_response(response)
...     scorer.submit()
```
success_message: Test for acidity detection implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_3'
>>> with patch('builtins.print') as mock_print:
...     new_pH_levels = [6.5, 6.8, 7.2, 4.5, 5.5, 6.0]
...     acidic_result = acidic_pH_detection(new_pH_levels, threshold_pH)
>>> if isinstance(acidic_result, tuple) and acidic_result == (3, 4.5):
...     score = 3
...     drexel_jupyter_logger.variable_logger_csv('3, 3', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 3, 'student_response': f'{acidic_result}, {type(acidic_result)}'}
>>> with patch('builtins.print') as mock_print:
...     scorer.add_response(response)
...     scorer.submit()
>>> assert isinstance(acidic_result, tuple), 'The function does not return a tuple.'
>>> assert acidic_result == (3, 4.5), 'The function does not return the correct index and pH value.'
```
success_message: Function correctly returns a numerical value.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_4'
>>> expected_message = 'No acidic solution was detected below the threshold of 5.0.'
>>> with patch('builtins.print') as mock_print:
...     threshold_pH = 5.0
...     out = acidic_pH_detection([6.5, 7.2, 7.8], threshold_pH)
...     mock_print.assert_called_once_with(expected_message)
>>> if mock_print.call_args[0][0] == expected_message and None == out:
...     score = 3
...     drexel_jupyter_logger.variable_logger_csv('3, 3', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 3, 'student_response': f'{mock_print.call_args[0][0]}, {out}'}
>>> with patch('builtins.print') as mock_print:
...     scorer.add_response(response)
...     scorer.submit()
>>> assert None == out, 'The function should return None for both index and value when no pH level below the threshold is detected.'
```
success_message: default acidity detection is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q2_5'
>>> with patch('builtins.print') as mock_print:
...     assert acidic_solution == (5, 4.8), 'The acidic solution index and pH value do not match the expected output.'
>>> if acidic_solution == (5, 4.8):
...     score = 4
...     drexel_jupyter_logger.variable_logger_csv('4, 4', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': 4, 'student_response': f'{acidic_solution}'}
>>> with patch('builtins.print') as mock_print:
...     scorer.add_response(response)
...     scorer.submit()
```
success_message: acidic_solution does not return the correct value.
