```python
# Initialize Otter
import otter
grader = otter.Notebook("quiz9-classes-practice.ipynb")
```

# ⌛️ Quiz 9 Practice - Classes

This quiz will evaluate your mastery of using classes in Python. Functions provide a way to isolate code that you want to use repeatedly, and they allow you to pass in data to the code and get data back out of the code.


## Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
import pkg_resources
from subprocess import call
import sys

package_name = 'ENGR131_Util_2024'
version = '0.1.11'
package_version = f'{package_name}=={version}'

try:
    # Check if the package and version are installed
    pkg_resources.require(package_version)
    print(f'{package_version} is already installed.')
except pkg_resources.DistributionNotFound:
    # If not installed, install the package
    print(f'{package_version} not found. Installing...')
    call([sys.executable, '-m', 'pip', 'install', package_version])
except pkg_resources.VersionConflict:
    # If a different version is installed, you can choose to upgrade/downgrade
    installed_packages = {dist.key: dist.version for dist in pkg_resources.working_set}
    installed_version = installed_packages.get(package_name.lower())
    print(f'{package_name} {installed_version} is installed, but {version} is required.')
    # Optionally, upgrade or downgrade the package to the required version
    call([sys.executable, '-m', 'pip', 'install', '--upgrade', package_version])
```

```python
# DO NOT MODIFY THIS CELL
from ENGR131_Util_2024 import cell_logger, StudentInfoForm, responses, upsert_to_json_file
# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', cell_logger)
responses["assignment"] = "quiz_9_practice"

StudentInfoForm(**responses)
```

## Question: Machine Learning Model Resource Estimator

### Background

In the field of machine learning and artificial intelligence, estimating the computational resources required for training a model is critical for efficiency and cost management. The total resource requirement includes both the static load (the base computational power and memory needed for the model architecture, datasets, and fixed algorithms) and the dynamic load (additional resources required for training iterations, data augmentation, and validation processes).

### Objectives

1. Implement a `ModelResourceEstimator` class to model the resource requirements for training a machine learning model.
2. Include methods within this class to calculate the static load, dynamic load, and total resource requirement.
3. Demonstrate the use of class initialization, basic math calculations, and functions calling other functions.

### Class to Implement

Implement a class `ModelResourceEstimator` with the following properties and methods:
   - Methods:
     - `__init__`: Initializes a new `ModelResourceEstimator` instance with the properties.
         - Properties:
           - `model_complexity` (a measure of model complexity, e.g., the number of parameters)
           - `data_size` (in gigabytes, GB, representing the size of the dataset used for training)
           - `iterations` (the number of training iterations)
           - `base_resource_per_iteration` (in gigaflops per iteration, representing the base computational resource usage per training iteration)
           - `memory_usage_static` (in gigabytes, GB, representing the static memory usage by the model and dataset)
     - `static_load`: Calculates and returns the static resource load for training the model.
       - The static load includes the base static memory usage.
     - `dynamic_load`: Calculates and returns the dynamic resource load for the model based on the number of iterations and base resource per iteration.
       - The dynamic load is calculated as the product of the number of iterations and the base resource usage per iteration.
     - `total_resource_requirement`: Calculates and returns the total resource requirement by summing its static and dynamic loads.
       - The total resource requirement is the sum of the static and dynamic loads.
     - `print_resource_requirements`: Prints the static, dynamic, and total resource requirements in the following format:
       - "Static Load: "{static_load} GB\nDynamic Load: {dynamic_load} GFLOPs\nTotal Resource Requirement: {total_resource_requirement} (GB for static load and GFLOPs for dynamic load)" where the values are extracted from the object rounded to 2 decimal places.
       - Note: you must use a **single print command**, \n is used to create a new line in the print statement.

Instantiate a `ModelResourceEstimator` object with the following properties:

   - Model complexity: High (not directly quantified here, but influences resource estimation)
   - Dataset size: 10 GB
   - Number of training iterations: 1000
   - Base computational resource per iteration: 5 gigaflops
   - Static memory usage: 2 GB


```python
# Your Class for a Machine Learning Model Resource Estimator goes Here
...

# Instantiate a ModelResourceEstimator object with the specified properties
...

# Print the static, dynamic, and total resource requirements
# you can uncomment this line for testing
# model_resource_estimator.print_resource_requirements()
```

```python
grader.check("q1-ML-Resource")
```


## Fin

Tests for q1-ML-Resource (points: 42)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> points_ = [6, 4, 5, 4, 4, 4, 4, 6, 5]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> import json
>>> student_data = {'first_name': responses['first_name'], 'last_name': responses['last_name'], 'drexel_id': responses['drexel_id'], 'drexel_email': responses['drexel_email']}
>>> with open('student_data.json', 'w') as json_file:
...     json.dump(student_data, json_file)
>>> scorer = submit_score()
>>> question_id = 'q1_1'
>>> max_score = 6
>>> score = 0
>>> for method in ['__init__', 'model_complexity', 'data_size', 'iterations', 'base_resource_per_iteration', 'memory_usage_static']:
...     if hasattr(model_resource_estimator, method):
...         score += 1
>>> score = int(np.ceil(score))
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': dir(model_resource_estimator)}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert hasattr(model_resource_estimator, '__init__'), '__init__ method is not implemented'
>>> assert hasattr(model_resource_estimator, 'model_complexity'), 'model_complexity method is not implemented'
>>> assert hasattr(model_resource_estimator, 'data_size'), 'data_size method is not implemented'
>>> assert hasattr(model_resource_estimator, 'iterations'), 'iterations method is not implemented'
>>> assert hasattr(model_resource_estimator, 'base_resource_per_iteration'), 'base_resource_per_iteration method is not implemented'
>>> assert hasattr(model_resource_estimator, 'memory_usage_static'), 'memory_usage_static method is not implemented'
```
success_message: All methods are implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q1_2'
>>> max_score = 2
>>> score = 0
>>> if len(ModelResourceEstimator.__init__.__code__.co_varnames) == 6:
...     score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': dir(model_resource_estimator)}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert len(ModelResourceEstimator.__init__.__code__.co_varnames) == 6, '__init__ method does not take the correct number of parameters'
```
success_message: Init takes the correct number of parameters
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> scorer = submit_score()
>>> question_id = 'q1_3'
>>> max_score = 5
>>> score = 0
>>> for method in ['model_complexity', 'data_size', 'iterations', 'base_resource_per_iteration', 'memory_usage_static']:
...     if hasattr(model_resource_estimator, method):
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': dir(model_resource_estimator)}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert hasattr(model_resource_estimator, 'model_complexity'), 'model_complexity attribute is not implemented'
>>> assert hasattr(model_resource_estimator, 'data_size'), 'data_size attribute is not implemented'
>>> assert hasattr(model_resource_estimator, 'iterations'), 'iterations attribute is not implemented'
>>> assert hasattr(model_resource_estimator, 'base_resource_per_iteration'), 'base_resource_per_iteration attribute is not implemented'
>>> assert hasattr(model_resource_estimator, 'memory_usage_static'), 'memory_usage_static attribute is not implemented'
```
success_message: All properties are implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_4'
>>> max_score = 4
>>> score = 0
>>> m1 = ModelResourceEstimator(model_complexity='High', data_size=10, iterations=1000, base_resource_per_iteration=5, memory_usage_static=2)
>>> test = m1.static_load() == 2
>>> statement = 'static_load method is not implemented correctly'
>>> value = m1.static_load()
>>> if test:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': value}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert test, statement
```
success_message: static_load method is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_5'
>>> max_score = 4
>>> score = 0
>>> m1 = ModelResourceEstimator(model_complexity='High', data_size=10, iterations=1000, base_resource_per_iteration=5, memory_usage_static=2)
>>> value = m1.static_load()
>>> test = m1.static_load() == 2
>>> statement = 'static load method is not implemented correctly'
>>> if test:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': value}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert test, statement
```
success_message: static load implemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_6'
>>> max_score = 4
>>> score = 0
>>> estimator = ModelResourceEstimator(model_complexity='High', data_size=10, iterations=1000, base_resource_per_iteration=5, memory_usage_static=2)
>>> value = estimator.dynamic_load()
>>> test = value == 5000
>>> statement = 'dynamic load method is not implemented correctly'
>>> if test:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': value}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert test, statement
```
success_message: dynamic load implemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> scorer = submit_score()
>>> question_id = 'q1_7'
>>> max_score = 4
>>> score = 0
>>> mre = ModelResourceEstimator(model_complexity='High', data_size=10, iterations=1000, base_resource_per_iteration=5, memory_usage_static=2)
>>> value = mre.total_resource_requirement()
>>> test = value == mre.static_load() + mre.dynamic_load()
>>> statement = 'Total resource requirement method is not implemented correctly'
>>> if test:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': value}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> assert test, statement
```
success_message: Total resource requirement implemented correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> scorer = submit_score()
>>> question_id = 'q1_9'
>>> max_score = 3
>>> expected_message = 'Static Load: 2.00 GB\nDynamic Load: 5000.00 GFLOPs\nTotal Resource Requirement: 5002.00 (GB for static load and GFLOPs for dynamic load)'
>>> with patch('builtins.print') as mock_print:
...     mre = ModelResourceEstimator(model_complexity='High', data_size=10, iterations=1000, base_resource_per_iteration=5, memory_usage_static=2)
...     mre.print_resource_requirements()
...     mock_print.assert_called_once_with(expected_message)
>>> if mock_print.call_args[0][0] == expected_message:
...     score = 6
...     drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
... else:
...     score = 0
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': f'{mock_print.call_args[0][0]}'}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
```
success_message: Input parameters are correctly defined.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import submit_score
>>> import numpy as np
>>> scorer = submit_score()
>>> question_id = 'q1_9'
>>> max_score = 5
>>> score = 0
>>> attributes = ['model_complexity', 'data_size', 'iterations', 'base_resource_per_iteration', 'memory_usage_static']
>>> values = ['High', 10, 1000, 5, 2]
>>> for (attribute, expected_value) in zip(attributes, values):
...     actual_value = eval(f'model_resource_estimator.{attribute}')
...     if attribute == 'model_complexity':
...         if actual_value == expected_value:
...             score += 1
...     elif actual_value == expected_value:
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': f'Model Complexity: {model_resource_estimator.model_complexity}, \n                                Data Size: {model_resource_estimator.data_size}, \n                                Iterations: {model_resource_estimator.iterations}, \n                                Base Resource per Iteration: {model_resource_estimator.base_resource_per_iteration}, \n                                Static Memory Usage: {model_resource_estimator.memory_usage_static}\n                                '}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> for (attribute, expected_value) in zip(attributes, values):
...     assert eval(f'model_resource_estimator.{attribute}') == expected_value
```
success_message: Input parameters are correctly defined.
