```python
# Initialize Otter
import otter
grader = otter.Notebook("lab3-image.ipynb")
```

# 🧪🖥 Lab 3: Image Processing

This lab explores basic image processing in Python.



## Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
# Please provide your first name, last name, Drexel ID, and Drexel email. Make sure these are provided as strings. "STRINGS ARE TEXT ENCLOSED IN QUOTATION MARKS."

# In the assignments you will see sections of code that you need to fill in that are marked with ... (three dots). Replace the ... with your code.
first_name = ...
last_name = ...
drexel_id = ...
drexel_email = ...

```

```python
grader.check("q0-Checking-Your-Name")
```

```python
import matplotlib.pyplot as plt
import numpy as np
```

**Loading and viewing Images** 

The matplotlib.pyplot module, which has been imported, can be used to load and view images. The following code loads an image called ```"bird.jpeg"```, scales it to have floating point pixel values from 0 to 1 (instead of 0 to 255), and stores it in a variable called ```img```. Make sure the provided ```bird.jpeg``` file is in the same directory as this notebook.

```python
img = plt.imread("bird.jpeg").astype("float") / 255
```

To view the image, we can use plt.imshow()

```python
plt.imshow(img)
```

This image is represented in Python as a 3 dimensional array with dimensions ($H$, $W$, $C$), where 

<center>

$H$ = height

$W$ = width

$C$ = number of color channels

</center>


To view the dimensions, we can print the img array's "shape" as below:

```python
print(img.shape)
```

From this, we can see that this image has a height of 854 pixels, a width of 1000 pixels, and 3 color channels. The 3 color channels represent the red, green, and blue components of an image. The color of each pixel in the full color image is determined by its RGB components. 

The following code lets us extract the three separate color channels from the image. We store these color channels – each with dimensions ($H$,$W$) – in variables `R`, `G`, and `B`.

```python
# For red, get the contents of the image at color channel index 0
R = img[:,:,0]

# For green, get the contents of the image at color channel index 1
G = img[:,:,1]

# For blue, get the contents of the image at color channel index 2
B = img[:,:,2]

```

We can view the 3 color channels individually like this:

```python
# create a plot containing 3 subplots
figure, plots = plt.subplots(ncols=3, nrows=1)

### plotting the red component
# start with an empty image array of zeros
r_img = np.zeros(img.shape)
# populate the image with the red color channel data from img
r_img[:,:,0] = R
# plot the red component image
plots[0].imshow(r_img)
# turn off the axis to make the plot look tidy
plots[0].set_axis_off()

### plotting the green component
# start with an empty image array of zeros
g_img = np.zeros(img.shape)
# populate the image with the green color channel data from img
g_img[:,:,1] = G
# plot the red component image
plots[1].imshow(g_img)
# turn off the axis to make the plot look tidy
plots[1].set_axis_off()

### plotting the blue component
# start with an empty image array of zeros
b_img = np.zeros(img.shape)
# populate the image with the blue color channel data from img
b_img[:,:,2] = B
# plot the red component image
plots[2].imshow(b_img)
# turn off the axis to make the plot look tidy
plots[2].set_axis_off()
```

**Task 1: To Black and White** 

Your task is to write a function that computes a black and white version of an image. A black and white (or grayscale) image has equal values for its $R$, $G$, and $B$ components. One way to create a black and white version of a color images, that we will use in this lab, is to average the three color components. The new components `newR`, `newG`, and `newB` would be computed as follows.

<center>

`newR` = `newG` = `newB` = $\frac{1}{3} $ (`R` $ + $ `G` $ + $ `B` $)$


</center>

Hint: Numpy arrays of equal size can be added, subtracted, multiplied, etc. as if the were numbers. For instance, two arrays of equal dimensions `a` and `b` can be added pixel-by-pixel with the single line:

<center>

`a` + `b`

</center>

Write python code to do the following:

* Define a function called ``to_bw`` which accepts as input a single argument, `I`, an array containing a color image.
* In the function, first extract the red, blue, and green color components of the image. (You may replicate how this was done above.)
* Compute the new components `newR`, `newG`, and `newB`, from `R`, `G`, and `B` using the equation above.
* Create an array of zeros called `bwI` and give it the same shape as the input array `I`.
* Populate the color channels of `bwI` to have the color components `newR`, `newG`, and `newB`.
* Return the grayscale image `bwI`

Hint: If you have a numpy array `X` and you want to create an array of zeros that has the same shape as `X`, do the following:
```np.zeros(X.shape)``` 

Your code replaces the prompt:  `...`


```python
...

```

```python
grader.check("task1-toBW")
```

View your black and white image:

```python
plt.imshow(to_bw(img));
```

**Task 2: To Sepia Tone** 

Sepia toning is a chemical process used in photography which gives an image a brownish coloring. Image editing software can approximate the effect of sepia toning using the given set of equations.

<center>

`newR` =  0.393 `R` + 0.769 `G` + 0.189 `B`

`newG` =  0.349 `R` + 0.686 `G` + 0.168 `B`

`newB` =  0.272 `R` + 0.534 `G` + 0.131 `B`

</center>

Finally, any resulting pixel values `> 1.0` get assigned a value of `1.0`.


Write python code to do the following:

* Define a function called ``to_sepia`` which accepts as input a single argument, `I`, an array containing a color image.
* As before, first extract the red, blue, and green color components of the image.
* Compute the new components `newR`, `newG`, and `newB`, from `R`, `G`, and `B` using the set of equations above.
* Create an array of zeros called `sepiaI` and give it the same shape as the input array `I`.
* Populate the color channels of `sepiaI` to have the color components `newR`, `newG`, and `newB`.
* For any values in the resulting image greater than $1$, replace them with the value $1$ 
* Return the sepia tone image `sepiaI`

Hint: Use `np.where(sepiaI > 1.0)` to get the indices in the sepia image where the value is greater than 1.


Your code replaces the prompt:  `...`


```python
...
```

```python
grader.check("task2-sepia")
```

View your sepia tone image:

```python
plt.imshow(to_sepia(img));
```

**Task 3: Cropping** 

It can be helpful to crop an image to make its size smaller and focus around an object of interest. Some AI image classification systems require training data that is cropped to a square region of interest.

In this problem your task is to crop the black and white image of the bird so that it is a square image centered around the bird. Specifically, you should crop a $601$ x $601$ square centered around the pixel with $(x,y)$ coordinates $(565,375)$. Note: The point $(0,0)$ is the top left of the image, and $y$ coordinates increase towards the bottom.

Write python code to do the following:

* Determine the desired range of $x$- and $y$-coordinates to crop
* Use array slicing to crop the black and white bird image and store it in a variable called `bw_bird_cropped`.

Hint: The code `I_crop = I[y1:y2+1,x1:x2+1]` crops image `I` between $x$ bounds [`x1`,`x2`] and $y$ bounds [`y1`,`y2`] and stores the resulting image in a variable called `I_crop`.


Your code replaces the prompt:  `...`


```python
...
```

```python
grader.check("task3-cropping")
```

Let's view the cropped image.

```python
plt.imshow(bw_bird_cropped);
```

**Task 4: Processing your own image** 

In this final task, you will load your own image and process it as if it were an input in an AI image classification network. You will choose an image relevant to your area of engineering, filter it to be black and white, and perform a square cropping around a region of interest. 

Write python code to do the following:

* Load an image of your choosing and store it in a variable called `my_img`
* Store a black and white version of the image in a varialbe called `bw_my_img`
* Store a cropped version of the black and white image in a variable called `bw_my_img_cropped`
* Your crop should be square of an appropriate size, centered around the region of interest.

Hint: You will need to upload your image so that it is in the same directory as this notebook.


Your code replaces the prompt:  `...`


```python
...
```

```python
grader.check("task4-own-image")
```

View your cropped image:

```python
plt.imshow(bw_my_img_cropped);
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/firQITtr) to navigate to the assignment on GitHub classroom.

**Important:** Make sure you submit your image file with your ipynb file.

If you need further instructions on submitting your assignment please look at Lab 1. 

## Viewing your score

Each `.ipynb` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

We have both public and hidden tests. You will be able to see the score of both tests, but not the specific details of why the test passed or failed. 

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q0-Checking-Your-Name (points: None)
```
>>> def get_local_part(email):
...     return email.split('@')[0]
>>> local_part_email = get_local_part(drexel_email)
>>> assert local_part_email == drexel_id, f"The DrexelID extracted from the email '{local_part_email}' does not match the Drexel ID '{drexel_id}'."
```
failure_message: It does not appear that you entered your information correctly. Please try again.
success_message: You have entered your information correctly.


Tests for task1-toBW (points: None)
```
>>> z = np.zeros((200, 200, 3))
>>> assert np.shape(z) == np.shape(to_bw(z))
```
success_message: Output is numpy array of the correct shape.
```
>>> z = np.zeros((10, 10, 3))
>>> z[0, 0, 0] = 0.3
>>> z[3, 3, 2] = 0.9
>>> z_bw = np.zeros((10, 10, 3))
>>> z_bw[0, 0] = np.array([0.1, 0.1, 0.1])
>>> z_bw[3, 3] = np.array([0.3, 0.3, 0.3])
>>> assert np.isclose(z_bw, to_bw(z)).all()
```
success_message: Correct output for input image 1.
```
>>> z = np.zeros((5, 4, 3))
>>> z[1, 3, 0] = 0.3
>>> z[1, 3, 1] = 0.5
>>> z[1, 3, 2] = 0.1
>>> z_bw = np.zeros((5, 4, 3))
>>> z_bw[1, 3] = np.array([0.3, 0.3, 0.3])
>>> assert np.isclose(z_bw, to_bw(z)).all()
```
success_message: Correct output for input image 2.
```
>>> z = np.ones((16, 3, 3))
>>> z[11, 1, 0] = 0
>>> z[11, 1, 2] = 0
>>> z_bw = np.ones((16, 3, 3))
>>> z_bw[11, 1] = np.array([0.333333333333, 0.333333333333, 0.333333333333])
>>> assert np.isclose(z_bw, to_bw(z)).all()
```
success_message: Correct output for input image 3.


Tests for task2-sepia (points: None)
```
>>> z = np.zeros((200, 200, 3))
>>> assert np.shape(z) == np.shape(to_sepia(z))
```
success_message: Output is numpy array of correct shape.
```
>>> z = np.zeros((4, 4, 3))
>>> z[0, 0, 0] = 0.3
>>> z[1, 2, 2] = 0.9
>>> z[1, 2, 1] = 0.96
>>> z[1, 2, 0] = 0.98
>>> z[3, 3, 2] = 0.9
>>> z_ans = np.array([0.1179, 0.1047, 0.0816, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 1.0, 1.0, 0.8971, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.1701, 0.1512, 0.1179])
>>> assert np.isclose(to_sepia(z).flatten(), z_ans).all()
```
success_message: Correct output for input image 1.
```
>>> z = np.zeros((6, 2, 3))
>>> z[0, 0, 0] = 0.8
>>> z[0, 0, 1] = 0.992
>>> z[3, 1, 1] = 0.96
>>> z[0, 1, 2] = 0.43
>>> z[0, 1, 0] = 0.56
>>> z_ans = np.array([1.0, 0.959712, 0.747328, 0.30135, 0.26768, 0.20865, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.73824, 0.65856, 0.51264, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0])
>>> assert np.isclose(to_sepia(z).flatten(), z_ans).all()
```
success_message: Correct output for input image 2.
```
>>> z = np.zeros((3, 4, 3))
>>> z[2, 2, 2] = 0.75
>>> z[1, 1, 1] = 0.5
>>> z[0, 0, 0] = 0.25
>>> z[2, 2, 1] = 0.85
>>> z[2, 2, 0] = 0.99
>>> z_ans = np.array([0.09825, 0.08725, 0.068, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.3845, 0.343, 0.267, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 1.0, 1.0, 0.82143, 0.0, 0.0, 0.0])
>>> assert np.isclose(to_sepia(z).flatten(), z_ans).all()
```
success_message: Correct output for input image 2.


Tests for task3-cropping (points: None)
```
>>> z = np.zeros((601, 601, 3))
>>> assert np.shape(bw_bird_cropped) == np.shape(z)
```
success_message: Output is a numpy array of correct shape.
```
>>> answer_patch = np.array([0.25359477, 0.2496732, 0.25359477, 0.25359477, 0.25620915, 0.25620915, 0.25751634, 0.26928105, 0.2627451, 0.22745098, 0.25620915, 0.23529412, 0.24836601, 0.27843137, 0.26666667, 0.24444444, 0.26405229, 0.13594771, 0.19084967, 0.29150327, 0.28104575, 0.25620915, 0.24444444, 0.2248366, 0.1751634, 0.18431373, 0.25620915, 0.50457516, 0.61045752, 0.74117647, 0.21437908, 0.16601307, 0.17777778, 0.2745098, 0.39869281, 0.55686275, 0.69281046, 0.78431373, 0.78823529, 0.79084967, 0.30457516, 0.43006536, 0.54117647, 0.67189542, 0.76862745, 0.79477124, 0.7869281, 0.75424837, 0.70849673, 0.65882353, 0.6745098, 0.78431373, 0.80653595, 0.79215686, 0.75816993, 0.69934641, 0.6496732, 0.64575163, 0.64444444, 0.65098039, 0.80522876, 0.76078431, 0.69803922, 0.64836601, 0.63267974, 0.64836601, 0.66666667, 0.62745098, 0.63137255, 0.62745098, 0.65359477, 0.65882353, 0.67712418, 0.67189542, 0.64836601, 0.62614379, 0.61176471, 0.63267974, 0.62745098, 0.63006536, 0.64313725, 0.6379085, 0.61960784, 0.60784314, 0.62091503, 0.63006536, 0.63006536, 0.61830065, 0.62222222, 0.63006536, 0.6248366, 0.62352941, 0.62352941, 0.62352941, 0.62352941, 0.62352941, 0.62352941, 0.62222222, 0.6248366, 0.62614379])
>>> assert np.isclose(bw_bird_cropped[100:110, 200:210, 0].flatten(), answer_patch, rtol=1.000001).all()
```
success_message: Output image is correctly cropped.


Tests for task4-own-image (points: None)
```
>>> assert bw_my_img_cropped.shape[0] == bw_my_img_cropped.shape[1]
```
success_message: Output has a square image shape.
```
>>> my_r = bw_my_img_cropped[:, :, 0]
>>> my_g = bw_my_img_cropped[:, :, 1]
>>> my_b = bw_my_img_cropped[:, :, 2]
>>> assert np.isclose(my_r, my_g).all() and np.isclose(my_g, my_b).all() and np.isclose(my_b, my_r).all()
```
success_message: Output image is black and white.
