```python
# Initialize Otter
import otter
grader = otter.Notebook("lab10-fitting.ipynb")
```

# 💻🧪 Lab 10: Combining Concepts to Build Interoperable Code for Plotting

This lab uses all of the concepts we have learned in python programming to make reusable building blocks for generating, plotting, and fitting data. 

As an engineer you will regularly be presented with noisy data and want to fit that data to a function. In this assignment, we will build some machinery and a coding schema to generate noisy data (since we do not have real data), plot data, and fit the results. We will build tools to simplify the visualization of the results. The code will be designed to be interoperable. If we were to have a new type of data or mathematical expression, we could reuse all the code we have written. 

## Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
import pkg_resources
from subprocess import call
import sys

package_name = 'ENGR131_Util_2024'
version = '1.0.0'
package_version = f'{package_name}=={version}'

try:
    # Check if the package and version are installed
    pkg_resources.require(package_version)
    print(f'{package_version} is already installed.')
except pkg_resources.DistributionNotFound:
    # If not installed, install the package
    print(f'{package_version} not found. Installing...')
    call([sys.executable, '-m', 'pip', 'install', package_version])
except pkg_resources.VersionConflict:
    # If a different version is installed, you can choose to upgrade/downgrade
    installed_packages = {dist.key: dist.version for dist in pkg_resources.working_set}
    installed_version = installed_packages.get(package_name.lower())
    print(f'{package_name} {installed_version} is installed, but {version} is required.')
    # Optionally, upgrade or downgrade the package to the required version
    call([sys.executable, '-m', 'pip', 'install', '--upgrade', package_version])
```

```python
from ENGR131_Util_2024 import responses, StudentInfoForm, cell_logger
from ENGR131_Util_2024 import submit_question, ResponseStore, ValidateLogFile

# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', cell_logger)
responses["assignment"] = "labs_10"

StudentInfoForm(**responses)
```

## 2. Importing Functions

Here we will import some packages that we need for this assignment

2.1 import `numpy` as `np`

2.2 import the submodule `matplotlib.pyplot` and assign it to `plt`

2.3 import from `scipy.optimize` the `curve_fit` function

```python
# 2.1 you need to import numpy as np
...

# 2.2 You need to import the submodule matplotlib.pyplot and assign it to plt
...

# 2.3 You need to import from scipy.optimize the curve_fit function
...
```

```python
grader.check("q2-Importing Functions")
```

## 3. Implementing a Class for MathExpressions

In Python and programming in general, it is common to have a base class that adds functionality to an object. For example, if you have a car that you are building, you might want to add blind spot detection. You can create the blind spot detection software and hardware and install it on multiple cars types. We will do the same with a math function. We will build a class that adds a methods to fit and evaluate a math expression. 

When we build this class we need it to be flexible to accept math functions with a flexible number of input parameters. We implement a fit function where we can input any parameters into the function (this is important for optimization required to fit data), and an evaluate method which uses the parameters set during initialization. 

Follow these steps:

3.1 Define a `class` `MathExpression`.

3.2 Build an initialization function. 
- The function should take a variable `func` which is a mathematical function
- The function should accept `**kwargs` 
    - The `kwargs` will be the parameters of the fitting function 
    - We want to save the `kwargs`, each key-value pair of `kwargs` where the key is the variable name, and `func` as an attribute of the object

3.3 Define a method `fit`.
- That accepts a required input `x`, and `**kwargs`
- Return the result when you call the objects method `func` that accepts a required input `x`, and `**kwargs`

3.4 Define an method for the `class` `evaluate`.
- Takes a required input `x`
- Returns the result when you call the objects method `func` that accepts a required input `x`, and the `**kwargs` values set at initialization.

```python
# 3.1 define a class MathExpression
...
    
    # 3.2 Use the built in initialization function
    # The function should take a variable func which is a mathematical function for generation and fitting (we will define this later)
    # The function should accept **kwargs
    # The kwargs will be the parameters of the fitting function 
    ...
        
        # use the .items() method to extract from kwargs a list of tuples with the key value pairs
        # Create a for loop that extracts the key value pairs from the kwargs
        #   make sure to unpack the tuple into the variables key and value
        # Inside the for loop, use the built in python method setattr to add an attribute to the MathExpression Object 
        #   for the key, use the variable key, 
        #   for the value, use the variable value
        ...
        
        # save the kwargs as an attribute of the object named kwargs
        ...
        
        # save the function as an attribute of the object named func
        ...
        
    # 3.3 for the class, define a method fit that accepts a required input x, and **kwargs
    #   x is the points to sample along the x axis
    #   **kwargs are new parameters to test
    #   Note: this is a requirement of the fitting function
    ...
        
        # return the result when you call the object's method func that accepts a required input x and **kwargs
        ...
        
    # 3.4 for the class, define a method evaluate that takes a required input x
    ...
       
        # return the result when you call the object's method func that accepts a required input x and the kwargs values set at initialization.
        ...
```

```python
grader.check("q3-Implement MathExpression")
```

## 4. Define a Linear Function

$$ Y = slope * x + intercept $$

We want to define a simple linear function. A linear model is one of the simplest fitting functions used in science and engineering. It says that an independent variable directly influences a dependent variable. 

4.1 Build a function `LinearFunction`
- `LinearFunction` should take three inputs: `x` for the independent variable, `m` for the slope, `b` for the intercept.
- Return the calculation of `y`, where y = slope * x + intercept. 
- Make sure to replace slope and intercept with the correct variables.

```python
# 4.1 build a function LinearFunction
# linear function should take two inputs m - for the slope, and b for the intercept
# return the calculation of y, where y = slope * x + intercept. 
# Make sure to replace slope and intercept with the correct variables.
...

```

```python
grader.check("q3 - Linear Function")
```

## 5. Plotting

One of the best ways to understand your code is doing what you expect is to visualize the results. We can do this by plotting the data. 

5.1 Instantiate an object `Linear_` using the `MathExpression` `Class`.
- The input `func` should be the `LinearFunction` you wrote
- The variables needed for kwargs must be provided as key value pairs. 
- In this case, those key value pairs are `m = 1`, and `b = -3`. 

5.2 Derive a linearly-spaced vector for the independent variable.
- Use the `np.linspace` function to make a linearly-spaced array from `0` to `10` with `100` steps

5.3. call the built-in `evaluate` method of the `LinearFunction` object over the range defined by `x`.
- assign the output to a variable `y`

5.4 Use the function `plt.plot` to plot `x` vs `y`.
- assign the plot to a variable called `linear_plot`

```python
# 5.1 Instantiate an object Linear_ using the MathExpression Class
# The variables needed for the Kwargs m = 1, and b = -3 need to be provided as key value pairs. 
# example my_function(x=1), this would provide a **kwargs key (x) value (1)

...

# 5.2 use the np.linspace function to make a linearly-spaced array from 0 to 10, with 100 steps
# save this to the variable x

...

# 5.3 call the built-in `evaluate` method of the LinearFunction object over the range defined by x
# assign the output to a variable y

...

# 5.4 use the function plt.plot to plot x vs y

...

# we do not need to add more labels to this plot. It was just for testing.
```

```python
grader.check("q5 - Plotting")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results are aaved. 

```python
submit_question()
```

# 6. A NoisyFunction

You are responsible for making a 3-4 minute VoiceThread submission that briefly explains each function of this class. *You are not responsible for writing any code or comments in this Task.*

We have written a function with several useful methods that you might want to use as an engineer. 

6.1 It inherits information from an object of type `MathExpression`

6.2 It can generate noisy data, which is to simulate real experimental data.

6.3 It can fit the data to an known function. 

6.4 It can fit, plot and format the results for quick visualization.

```python
# build a class Noisy Function that Inherits MathExpression
# This is simply done writing class NoisyFunction(MathExpression)
class NoisyFunction(MathExpression):
    
    # This class does not require any specific initialization variables, but it will require that we pass the inherited variable through.
    # we have done this for you 
    # you should not touch this code. If you do, here it is
    # def __init__(self, *args, **kwargs):
    #    super().__init__(*args, **kwargs)
    def __init__(self, *args, **kwargs):
        super().__init__(*args, **kwargs)

    # build a function generate noisy data with two required inputs
    #   x_range - a list of the minimum and max ranges for to generate the data
    #   noise_amplitude - a float from which sets the magnitude of the noise
    def generate_noisy_data(self, x_range, noise_amplitude):
        # use np.linspace to generate x data along the range
        # generate 100 points

        x = np.linspace(x_range[0], x_range[1], 100)
        
        # call the evaluate method of the base function object to return the true values

        y_true = self.evaluate(x)
        
        # generate the noise vector by multiplying the noise amplitude by an random noise vector of length x  
        # the random noise vector should be generated with np.random.randn

        noise = noise_amplitude * np.random.randn(len(x))
        
        # add the noise to the y_true value
        # save it in a new variable y

        y = y_true + noise
        
        # return x, y, and y_true

        return x, y, y_true

    # build a function fit_data that takes 2 inputs x and y
    def fit_data(self, x, y):
        # call the curvefit function
        # provide the inherited pointed to the function
        # For the initial guess provide the true values they are contained within the saved variable kwargs
        # You can get the values of a dictionary using the built-in method .values()
        # P0 takes a list. You need to convert the values to a list.
        # Curvefit should return two values assign them to popt and pcov. 
        # popt is the fit results

        popt, pcov = curve_fit(self.func, x, y, p0=list(self.kwargs.values()))
        
        # Return popt and pcov

        return popt, pcov

    # build a function plot_fit_results that plots the raw and fit results
    # This function should take two input parameters:
    #   x_range - a tuple defining the x range to plot
    #   noise_amplitude - a float used to control the magnitude of the noise
    def plot_fit_results(self, x_range, noise_amplitude):
    

            # use the subplots function in pyplot to create a multiplot figure with 2 columns and 1 row.
            # set the figure size = 10, 5 - figsize takes a tuple
            # assign the output to two variables fig and axs
    
            fig, axs = plt.subplots(1, 2, figsize=(10, 5))
            

            # make a list of strings called labels 
            # The 0th index should be "True Fit"
            # The 1st index should be "Noisy Fit"
    
            labels = ['True Fit', 'Noisy Fit']
            

            # call the generate_noisy_data function from this class
            # save the returned variables to x, y, and y_true
    
            x, y, y_true = self.generate_noisy_data(x_range, noise_amplitude)
            

            # make a look that loops over the axs with enumerate.
            # set the iterator equal to the variable i
            # set the value equal to the variable ax
    
            for i, ax in enumerate(axs):
                
    
                # create a scatterplot on the axis object 
                # plot x vs y
                # make the marker size = 5 (with the `s` tag)
                # add a label 'Raw Data' (with the `label` tag)
        
                ax.scatter(x, y, s=5, label='Raw Data')
                
    
                # write an if statement that is true if the label of the graph is "True Fit". This is ordered by the list labels you created.
        
                if labels[i] == 'True Fit':
                    
        
                    # if the statement was true on the axis object plot
                    # x vs y true
                    # make the line red using the 'r' flag
                    # add a label as indexed from the list
            
                    ax.plot(x, y_true, 'r', label=labels[i])
                
    
                # add an elif statement that is true if the label of the graph is "Noisy Fit" this is ordered by the list labels you created.
        
                elif labels[i] == 'Noisy Fit':
                    
        
                    # call the method in the class fit_data
                    # This should fit the x and y data
                    # have it return the variables popt and pcov
            
                    popt, pcov = self.fit_data(x, y)
                    
        
                    # use the predicted popt and the inherited function to generate the y values for the fit results
                    # make sure to unpack the popt list with the * operator
            
                    y_noisy_fit = self.func(x, *popt)
                    
        
                    # on the axis object plot
                    # x vs y_noisy_fit
                    # make the line blue using the 'b' flag
                    # add a label as indexed from the list
            
                    ax.plot(x, y_noisy_fit, 'b', label=labels[i])
                 
    
                # else passes to the next loop
            
                else: 
                    pass
            
    
                # use the set_xlabel method to set the x-label of the graph to 'x'
                # use the set_ylabel method to set the y-label of the graph to 'y'
        
                ax.set_xlabel('x')
                ax.set_ylabel('y')
                
    
                # call the legend method of the axis object to show the legend
        
                ax.legend()
                               
            return fig

```

## 7. Testing our Class

In coding it is always a good idea to run all parts of your code to check it is performing as expected. 

7.1 Create a `NoisyFunction` instance for a linear function
- Assign this object to the variable `linear_func`
- The input `func` should be the `LinearFunction` you wrote
- The slope should be `3` and the intercept should be `4`

7.2 Call the proper method of `linear_func` to generate noisy data over the range from 0- $2\pi$
- the noise amplitude is `1`
- save the results to `x`, `y`, and `y_true`

7.3 Fit the `x` and `y` data using the proper method of `linear_func`
- save the fit results to `popt`, `pcov`
- (these names represent the optimal parameters (`popt`) and the parameters' covariance (`pcov`))

7.4 Use the built-in method to plot the fit result
- set the range equal to `0` to `16`
- set the noise amplitude equal to `0.3`
- Save this to the variable plot


```python
seed = np.random.seed(42)

# 7.1 Create a NoisyFunction instance for a linear function
# Assign this object to the variable linear_func
# The slope should be 3 and the intercept should be 4
linear_func = ...

# 7.2 call the generate_noisy_data function of the object the linear_func to produce noisy data over the range from 0-2pi
#   use a tuple to indicate the range from 0-2pi
#   the noise amplitude is 1
#   save the results to x, y, and y_true
x, y, y_true  = ...

# 7.3 Fit the data
# save the fit results to popt, pcov
popt, pcov = ...

# 7.4 Use the built-in method to plot the fit result
# set the range equal to 0 to 16
# set the noise amplitude equal to 0.3
# Save this to the variable plot
plot = ...

```

```python
grader.check("q7 - plotting and fitting")
```

## 8. Defining new math functions

Now we can view the flexibility (interoperability) of our code. 

Define two new functions. Follow the logic of the way you created `LinearFunction`.

8.1 `SineFunction`
   
$$g(x) = A (sin(2\pi f x+\phi))$$

where $A$ represents amplitude, $f$ represents the frequency, and $\phi$ represents the phase.

8.2 `ExponentialFunction`

$$h(x) = ae^{bx}$$
   

```python
# 8.1 Write the SineFunction here
# the inputs should be x, amplitude, frequency, phase
...

# 8.1 Write the ExponentialFunction here
# the inputs should be x, a, b
...
```

```python
grader.check("q8-new functions")
```

# 9 Using our class with the sine function

Now we can use nearly the same code from Task 7 to test our `SineFunction`.

9.1 Create a `NoisyFunction` instance for the sine function
- Assign this object to the variable `sine_func`
- The input `func` should be the `SineFunction` you wrote
- The amplitude should be 2, the frequency should be .2, and the phase should be .5

9.2 Apply the `generate_noisy_data` method of `sine_func` to produce noisy data over the range from 0-6 $\pi$
- use a tuple to indicate the range from 0-6 $\pi$
- the noise amplitude is 1.5
- save the results to `x`, `y`, and `y_true`

9.3 Fit the `x` and `y` data using the proper method of `sine_func`
- save the fit results to `popt`, `pcov`

9.4 Use the built-in method to plot the fit results
- set the range equal to 0 to 6 $\pi$
- set the noise amplitude equal to 1.5
- Save this to the variable `plot`

```python
seed = np.random.seed(42)

# 9.1 Create a NoisyFunction instance for a sine function
# Assign this object to the variable sine_func
# The amplitude should be 2, the frequency should be .2, and the phase should be .5
sine_func = ...

# 9.2 evaluate the sine_func to produce noisy data over the range from 0-6pi
# the noise amplitude is 1.5
# save the results to x, y, and y_true
x, y, y_true  = ...

# 9.3 Fit the data
# save the fit results to popt, pcov
popt, pcov = ...

# 9.4 Use the built-in method to plot the fit results
# set the range equal to 0 to 6 pi
# set the noise amplitude equal to 1.5
# Save this to the variable plot
plot = ...
```

```python
grader.check("q9-using the class with new Sine functions")
```

# 10 Using our class with the Exponential function

Now we can use nearly the exact same code to test our Exponential function. We wrote one code that we can use over and over again in different ways!

10.1 Create a `NoisyFunction` instance for a exponential function
- Assign this object to the variable `exp_func`
- The input `func` should be the `ExponentialFunction` you wrote
- `a` should be 2, and `b` should be .2

10.2 Call the proper method of `exp_func` to generate noisy data over the range from 0-6 $\pi$
- use a tuple to indicate the range from 0-6 $\pi$
- the noise amplitude is 1.5
- save the results to `x`, `y`, and `y_true`

10.3 Fit the `x` and `y` data using the proper method of `exp_func`
- save the fit results to `popt`, `pcov`

10.4 Use the built-in method to plot the fit results
- set the range equal to 0 to 6 $\pi$
- set the noise amplitude equal to 1.5
- save this to the variable `plot`

```python
seed = np.random.seed(42)

# 10.1 Create a NoisyFunction instance for a exponential function
# Assign this object to the variable exp_func
# The a should be 2, and b should be .2
exp_func = ...

# 10.2 evaluate the exp_func to produce noisy data over the range from 0-6pi
# the noise amplitude is 1.5
# save the results to x, y, and y_true
x, y, y_true  = ...

# 10.3 Fit the data
# save the fit results to popt, pcov
popt, pcov = ...

# 10.4 Use the built-in method to plot the fit results
# set the range equal to 0 to 6 pi
# set the noise amplitude equal to 1.5
# Save this to the variable plot
plot = ...
```

```python
grader.check("q10-using the class with new exponential functions")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results are aaved. 

```python
submit_question()
ValidateLogFile("./output.log", responses["assignment"], 0)
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/xlac6lYB) to navigate to the assignment on GitHub classroom.

If you need further instructions on submitting your assignment please look at Lab 1. 

## Viewing your score

**It is your responsibility to ensure that your grade report shows correctly. We can only provide corrections to grades if a grading error is determined. If you do not receive a grade report your grade has not been recorded. It is your responsibility either resubmit the assignment correctly or contact the instructors before the assignment due date.**

Each `.ipynb` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

We have both public and hidden tests. You will be able to see the score of both tests, but not the specific details of why the test passed or failed. 

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q10-using the class with new exponential functions (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> points_ = [1, 0, 1, 4, 5]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_1'
>>> max_score = 1
>>> score = 0
>>> if isinstance(exp_func, NoisyFunction) and exp_func.a == 2 and (exp_func.b == 0.2) and (exp_func.func == ExponentialFunction):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert isinstance(exp_func, NoisyFunction)
>>> assert exp_func.a == 2
>>> assert exp_func.b == 0.2
>>> assert exp_func.func == ExponentialFunction
```
failure_message: exp_func implemented incorrectly
success_message: exp_func implemented correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_2'
>>> max_score = 0
>>> score = 0
>>> if seed == np.random.seed(42):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert seed == np.random.seed(42)
```
failure_message: You might have deleted the line that sets the random seed. Please add it at the top of the cell. You will fail the autograder if you do not place the seed in the top of the cell.
success_message: Your seed is still correct
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_3'
>>> max_score = 1
>>> score = 0
>>> if np.isclose(sum([x, y, y_true])[4], 5.068497807066033, rtol=0.01) and np.isclose(np.mean(sum([x, y, y_true])[4::3]), 54.848917146243046, rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(sum([x, y, y_true])[4], 5.068497807066033, rtol=0.01)
>>> assert np.isclose(np.mean(sum([x, y, y_true])[4::3]), 54.848917146243046, rtol=0.01)
```
failure_message: x, y, y_true return incorrectly
success_message: x, y, y_true return correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_4'
>>> max_score = 4
>>> score = 0
>>> if np.isclose(popt[0], 1.9635667563774726, rtol=0.01) and np.isclose(popt[1], 0.20098217114795966, rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(popt[0], 1.9635667563774726, rtol=0.01)
>>> assert np.isclose(popt[1], 0.20098217114795966, rtol=0.01)
```
failure_message: Fitting function incorrectly implemented
success_message: Fitting function correctly implemented
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_5'
>>> max_score = 5
>>> score = 0
>>> if np.isclose(np.sum(plot.axes[0].get_lines()[0].get_xydata()), 3212.76954915798, rtol=0.01) and np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 3220.812931722522, rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(np.sum(plot.axes[0].get_lines()[0].get_xydata()), 3212.76954915798, rtol=0.01)
>>> assert np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 3220.812931722522, rtol=0.01)
```
failure_message: plot incorrectly implemented
success_message: plot implemented correctly


Tests for q2-Importing Functions (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> points_ = [1, 1, 1]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q1_1'
>>> max_score = 1
>>> score = 0
>>> if np.__version__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.__version__ is not None
```
failure_message: numpy random incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q1_2'
>>> max_score = 1
>>> score = 0
>>> if plt.__package__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert plt.__package__ is not None
```
failure_message: plt random incorrectly implemented.
success_message: plt correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q1_3'
>>> max_score = 1
>>> score = 0
>>> if curve_fit.__name__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert curve_fit.__name__ is not None
```
failure_message: plt random incorrectly implemented.
success_message: plt correctly implemented.


Tests for q3 - Linear Function (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> from inspect import isfunction
>>> points_ = [1, 1, 1]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q3_1'
>>> max_score = 1
>>> score = 0
>>> if isfunction(LinearFunction):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert isfunction(LinearFunction)
```
failure_message: LinearFunction is not a function
success_message: LinearFunction is a function
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import inspect
>>> signature = inspect.signature(LinearFunction)
>>> num_parameters = len(signature.parameters)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q3_2'
>>> max_score = 1
>>> score = 0
>>> if num_parameters == 2:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert num_parameters == 3
```
failure_message: LinearFunction function has the incorrect inputs.
success_message: LinearFunction function has the correct inputs.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> x = np.array([2, 3, 4, 5, 8, 10])
>>> out = LinearFunction(x, 7, 4)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q3_3'
>>> max_score = 1
>>> score = 0
>>> if all(out == [18, 25, 32, 39, 60, 74]):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert all(out == [18, 25, 32, 39, 60, 74])
```
failure_message: LinearFunction function implemented incorrect inputs.
success_message: LinearFunction function implemented correctly


Tests for q3-Implement MathExpression (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> points_ = [1, 1, 2, 1, 1, 3, 1, 3]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q2_1'
>>> max_score = 1
>>> score = 0
>>> if isinstance(MathExpression, type):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert isinstance(MathExpression, type)
```
failure_message: MathExpression is not a Class
success_message: MathExpression is a Class
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q2_2'
>>> max_score = 1
>>> score = 0
>>> if '__init__' in dir(MathExpression):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert '__init__' in dir(MathExpression)
```
failure_message: __init__ not implemented.
success_message: __init__ implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> out = MathExpression('a', b=1, test='test')
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q2_3'
>>> max_score = 2
>>> score = 0
>>> if out.kwargs['b'] == 1 and out.kwargs['test'] == 'test':
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert out.kwargs['b'] == 1
>>> assert out.kwargs['test'] == 'test'
```
failure_message: MathExpression inputs kwargs not correctly implemented.
success_message: MathExpression inputs kwargs correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> out = MathExpression(np.sum)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q2_4'
>>> max_score = 1
>>> score = 0
>>> if out.func is np.sum:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert out.func is np.sum
```
failure_message: MathExpression function input not correctly implemented.
success_message: MathExpression function input correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q2_5'
>>> max_score = 1
>>> score = 0
>>> if 'fit' in dir(MathExpression):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert 'fit' in dir(MathExpression)
```
failure_message: fit is not implemented.
success_message: fit is implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> 
>>> def test(x, **kwargs):
...     return (x, kwargs)
>>> out = MathExpression(test)
>>> (a, b) = out.fit(2, a=2, b=3)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q2_6'
>>> max_score = 3
>>> score = 0
>>> if a == 2 and b == {'a': 2, 'b': 3}:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert a == 2
>>> assert b == {'a': 2, 'b': 3}
```
failure_message: fit is not implemented correctly.
success_message: fit is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q2_7'
>>> max_score = 1
>>> score = 0
>>> if 'evaluate' in dir(MathExpression):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert 'evaluate' in dir(MathExpression)
```
failure_message: evaluate is not implemented.
success_message: evaluate is implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> 
>>> def test(x, **kwargs):
...     return (x, kwargs)
>>> out = MathExpression(test, a=5, b=3)
>>> (a, b) = out.evaluate(2)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q2_8'
>>> max_score = 3
>>> score = 0
>>> if a == 2 and b == {'a': 5, 'b': 3}:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert a == 2
>>> assert b == {'a': 5, 'b': 3}
```
failure_message: evaluate is not implemented correctly.
success_message: evaluate is implemented correctly.


Tests for q5 - Plotting (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> points_ = [1, 1, 1, 3]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q4_1'
>>> max_score = 2
>>> score = 0
>>> if '__init__' in dir(MathExpression):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert isinstance(Linear_, MathExpression)
>>> assert Linear_.m == 1
>>> assert Linear_.b == -3
```
failure_message: MathExpression is not instantiated correctly
success_message: MathExpression is instantiated correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q4_2'
>>> max_score = 1
>>> score = 0
>>> if np.isclose(x[0], 0) and np.isclose(x[-1], 10) and (x.__len__() == 100):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(sum(x[::4]), 121.212121, rtol=0.01)
>>> assert x.__len__() == 100
```
failure_message: X vector is incorrect
success_message: X vector is correct
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q4_3'
>>> max_score = 1
>>> score = 0
>>> if sum(y) == 200:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert sum(y) == 200
```
failure_message: evaluate computed incorrectly
success_message: evaluate computed correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q4_4'
>>> max_score = 3
>>> score = 0
>>> if 700 == np.sum(linear_plot[0].get_data()):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert 700 == np.sum(linear_plot[0].get_data())
```
failure_message: The plot is incorrect
success_message: The plot is correct


Tests for q7 - plotting and fitting (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> points_ = [1, 0, 1, 4]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_1'
>>> max_score = 1
>>> score = 0
>>> if isinstance(linear_func, NoisyFunction) and linear_func.m == 3 and (linear_func.b == 4) and (linear_func.func == LinearFunction):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert isinstance(linear_func, NoisyFunction)
>>> assert linear_func.m == 3
>>> assert linear_func.b == 4
>>> assert linear_func.func == LinearFunction
```
failure_message: linear_func implemented incorrectly
success_message: linear_func implemented correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_2'
>>> max_score = 0
>>> score = 0
>>> if seed == np.random.seed(42):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert seed == np.random.seed(42)
```
failure_message: You might have deleted the line that sets the random seed. Please add it at the top of the cell. You will fail the autograder if you do not place the seed in the top of the cell.
success_message: Your seed is still correct
```
>>> seed = np.random.seed(42)
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_3'
>>> max_score = 1
>>> score = 0
>>> if np.isclose(sum([x, y, y_true])[4], 9.542909136398162, rtol=0.01) and np.isclose(np.mean(sum([x, y, y_true])[4::3]), 30.365526866857397, rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(sum([x, y, y_true])[4], 9.542909136398162, rtol=0.01)
>>> assert np.isclose(np.mean(sum([x, y, y_true])[4::3]), 30.365526866857397, rtol=0.01)
```
failure_message: x, y, y_true return incorrectly
success_message: x, y, y_true return correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_4'
>>> max_score = 4
>>> score = 0
>>> if np.isclose(popt[0], 3.0219526649250503, rtol=0.01) and np.isclose(popt[1], 3.8271871517503655, rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(popt[0], 3.0219526649250503, rtol=0.01)
>>> assert np.isclose(popt[1], 3.8271871517503655, rtol=0.01)
```
failure_message: Fitting function incorrectly implemented
success_message: Fitting function correctly implemented
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_5'
>>> max_score = 2
>>> score = 0
>>> if '__init__' in dir(MathExpression):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.sum(plot.axes[0].get_lines()[0].get_xydata()) == 3600
>>> assert np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 3600.6691376113495, rtol=0.01)
```
failure_message: plot incorrectly implemented
success_message: plot implemented correctly


Tests for q8-new functions (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import inspect
>>> signature = inspect.signature(SineFunction)
>>> num_parameters = len(signature.parameters)
>>> seed = np.random.seed(42)
>>> points_ = [3, 3]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q6_1'
>>> max_score = 3
>>> score = 0
>>> if num_parameters == 4 and np.allclose([-1.51360499, -0.77557439, -0.456128], SineFunction(np.array([0.0, 0.1, 0.3]), 2, 3, 4), rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert num_parameters == 4
>>> assert np.allclose([-1.51360499, -0.77557439, -0.456128], SineFunction(np.array([0.0, 0.1, 0.3]), 2, 3, 4), rtol=0.01)
```
failure_message: SineFunction incorrectly implemented
success_message: SineFunction implemented correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import inspect
>>> signature = inspect.signature(ExponentialFunction)
>>> num_parameters = len(signature.parameters)
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q6_2'
>>> max_score = 3
>>> score = 0
>>> if num_parameters == 3 and np.allclose([2.0, 2.69971762, 4.91920622], ExponentialFunction(np.array([0.0, 0.1, 0.3]), 2, 3), rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert num_parameters == 3
>>> assert np.allclose([2.0, 2.69971762, 4.91920622], ExponentialFunction(np.array([0.0, 0.1, 0.3]), 2, 3), rtol=0.01)
```
failure_message: SineFunction incorrectly implemented
success_message: SineFunction implemented correctly


Tests for q9-using the class with new Sine functions (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> points_ = [1, 0, 1, 4, 5]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q1_{i + 1}')
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_1'
>>> max_score = 1
>>> score = 0
>>> if isinstance(sine_func, NoisyFunction) and sine_func.amplitude == 2 and (sine_func.frequency == 0.2) and (sine_func.phase == 0.5) and (sine_func.func == SineFunction):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert isinstance(sine_func, NoisyFunction)
>>> assert sine_func.amplitude == 2
>>> assert sine_func.frequency == 0.2
>>> assert sine_func.phase == 0.5
>>> assert sine_func.func == SineFunction
```
failure_message: sine_func implemented incorrectly
success_message: sine_func implemented correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_2'
>>> max_score = 0
>>> score = 0
>>> if seed == np.random.seed(42):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert seed == np.random.seed(42)
```
failure_message: You might have deleted the line that sets the random seed. Please add it at the top of the cell. You will fail the autograder if you do not place the seed in the top of the cell.
success_message: Your seed is still correct
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_3'
>>> max_score = 1
>>> score = 0
>>> if np.isclose(sum([x, y, y_true])[4], 4.384520747622161, rtol=0.01) and np.isclose(np.mean(sum([x, y, y_true])[4::3]), 9.50518884876772, rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(sum([x, y, y_true])[4], 4.384520747622161, rtol=0.01)
>>> assert np.isclose(np.mean(sum([x, y, y_true])[4::3]), 9.50518884876772, rtol=0.01)
```
failure_message: x, y, y_true return incorrectly
success_message: x, y, y_true return correctly
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_4'
>>> max_score = 4
>>> score = 0
>>> if np.isclose(popt[0], 2.2550781375831015, rtol=0.01) and np.isclose(popt[1], 0.19762745754501546, rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(popt[0], 2.2550781375831015, rtol=0.01)
>>> assert np.isclose(popt[1], 0.19762745754501546, rtol=0.01)
```
failure_message: Fitting function incorrectly implemented
success_message: Fitting function correctly implemented
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> seed = np.random.seed(42)
>>> responder = ResponseStore()
>>> responder.create_file()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_5'
>>> max_score = 5
>>> score = 0
>>> if np.isclose(np.sum(plot.axes[0].get_lines()[0].get_xydata()), 944.5788167212941, rtol=0.01) and np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 943.8768017567118, rtol=0.01):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': ''}
>>> responder.add_response(response)
>>> assert np.isclose(np.sum(plot.axes[0].get_lines()[0].get_xydata()), 944.5788167212941, rtol=0.01)
>>> assert np.isclose(np.sum(plot.axes[1].get_lines()[0].get_xydata()), 943.8768017567118, rtol=0.01)
```
failure_message: plot incorrectly implemented
success_message: plot implemented correctly
