```python
# Initialize Otter
import otter
grader = otter.Notebook("hw8-strings.ipynb")
```

# Homework 8

This assignment includes three problems on the topic of strings.



## Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
# Please provide your first name, last name, Drexel ID, and Drexel email. Make sure these are provided as strings. "STRINGS ARE TEXT ENCLOSED IN QUOTATION MARKS."

# In the assignments you will see sections of code that you need to fill in that are marked with ... (three dots). Replace the ... with your code.
first_name = ...
last_name = ...
drexel_id = ...
drexel_email = ...

```

```python
grader.check("q0-Checking-Your-Name")
```

**Question 1: Remove non alpha characters** 

In this lab you will write a function which takes in a string and removes all of the non-alphabetic characters from it.

Write python code to do the following:

* Define a function called `remove_non_alpha` that takes an input string `s`
* Loop through each character in the string, adding it to an output string if it is an alphabetic character.
* Return the final string.


Your code replaces the prompt:  `...`


```python
...
```

```python
grader.check("q1-nonalpha")
```

**Question 2: Palindrome Detector** 

A palindrome is a word or a phrase that is the same when read both forward and backward. Examples are: "bob," "sees," or "never odd or even" (ignoring spaces). Write a program whose input is a word or phrase, and that outputs true if the input is a palindrome, otherwise false. In determining if a string is a palindrome, you should treat upper vs. lower case versions of characters as the same letter, and you may ignore all spaces. 

Write python code to do the following:

* Define a function called is_palindrome that takes in one input parameter `x`.
* Convert the string `x` to be all lower case.
* Remove all spaces from string `x`. (Hint: use the string.`replace` method)
* Loop through the characters of `x` to determine if it is palindromic

Your code replaces the prompt:  `...`

```python
...

```

```python
grader.check("q2-palindrome")
```

**Question 3: Count characters** 

Your task is to write a function that takes in a character and a string and retruns the count of how many times that character appears in the string.

Write python code to do the following:

* Define a function called ``count_characters`` which takes two arguments, a character `c`, and a string `s`
* Your function should loop through the string and character by character, counting every time the character `c` is encountered
* Return the total count, or `0` if the character `c` is not part of string `s`


Your code replaces the prompt:  `...`


```python
...
```

```python
grader.check("q3-countchars")
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/Zglxmfb4) to navigate to the assignment on GitHub classroom.

If you need further instructions on submitting your assignment please look at Lab 1. 

## Viewing your score

**It is your responsibility to ensure that your grade report shows correctly. We can only provide corrections to grades if a grading error is determined. If you do not receive a grade report your grade has not been recorded. It is your responsibility either resubmit the assignment correctly or contact the instructors before the assignment due date.**

Each `.ipynb` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

We have both public and hidden tests. You will be able to see the score of both tests, but not the specific details of why the test passed or failed. 

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q0-Checking-Your-Name (points: None)
```
>>> def get_local_part(email):
...     return email.split('@')[0]
>>> local_part_email = get_local_part(drexel_email)
>>> assert local_part_email == drexel_id, f"The DrexelID extracted from the email '{local_part_email}' does not match the Drexel ID '{drexel_id}'."
```
failure_message: It does not appear that you entered your information correctly. Please try again.
success_message: You have entered your information correctly.


Tests for q1-nonalpha (points: 8)
```
>>> assert remove_non_alpha('Test1') == 'Test'
```
success_message: remove_non_alpha("Test1") is correct.
```
>>> assert remove_non_alpha('2Te2st2') == 'Test'
```
success_message: remove_non_alpha("2Te2st2") is correct.
```
>>> assert remove_non_alpha('abc123---xyz') == 'abcxyz'
```
success_message: remove_non_alpha("abc123---xyz") is correct.
```
>>> assert remove_non_alpha('42!') == ''
```
success_message: remove_non_alpha("42!") is correct.
```
>>> assert remove_non_alpha('same') == 'same'
```
success_message: remove_non_alpha("same") is correct.


Tests for q2-palindrome (points: None)
```
>>> assert is_palindrome('0')
```
success_message: is_palindrome("0") is correct.
```
>>> assert is_palindrome('bob')
```
success_message: is_palindrome("bob") is correct.
```
>>> assert is_palindrome('RACE car')
```
success_message: is_palindrome("RACE car") is correct.
```
>>> assert not is_palindrome('race! car!')
```
success_message: is_palindrome("race! car!") is correct.
```
>>> assert not is_palindrome('ABC-DEF')
```
success_message: is_palindrome("ABC-DEF") is correct.


Tests for q3-countchars (points: 5)
```
>>> assert count_characters('h', 'hello') == 1
```
success_message: count_characters("h","hello") is correct.
```
>>> assert count_characters('l', 'hello') == 2
```
success_message: count_characters("l","hello") is correct.
```
>>> assert count_characters('a', 'testing') == 0
```
success_message: count_characters("a","testing") is correct.
```
>>> assert count_characters('b', 'aaabbbcccbbbaaa') == 6
```
success_message: count_characters("b","aaabbbcccbbbaaa") is correct.
```
>>> assert count_characters('b', 'a b c') == 1
```
success_message: count_characters("b","a b c") is correct.
