# ⌛️ Quiz 6 - Using Functions

This quiz will evaluate your mastery of using functions in Python. Functions provide a way to isolate code that you want to use repeatedly, and they allow you to pass in data to the code and get data back out of the code.


## Instructions

You will click on the following link to start your quiz. Do not start the quiz before your lab section. We have checks in place to ensure that you do not start the quiz early. If you do, you will receive a 0 on the quiz. Furthermore, it could be considered an academic integrity violation.

[CLICK HERE TO START YOUR QUIZ IN LAB](https://online.wiseattend.com/student/testLink?c=ENGR131&k=0464) 



## Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
from ENGR131_Util_2024 import engr131_midterm

# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', engr131_midterm)

# Please provide your first name, last name, Drexel ID, and Drexel email. Make sure these are provided as strings. "STRINGS ARE TEXT ENCLOSED IN QUOTATION MARKS."

# In the assignments you will see sections of code that you need to fill in that are marked with ... (three dots). Replace the ... with your code.
first_name = ...
last_name = ...
drexel_id = ...
drexel_email = ...

```

**Question 1: Stress on a Beam**

You are a structural engineer responsible for selecting a beam to support a load. You know the beam is subjected to a uniform load. It is your job to create a function that will calculate the minimum beam design specifications to ensure that the beam will not yield under the expected load.

The stress, $\sigma$, on a beam subjected to a uniform load is given by the equation:

$$ \sigma = \frac{M y}{I} $$

where:
- $M$ is the moment of the force applied to the beam (in Newton-meters, Nm),
- $y$ is the distance from the neutral axis to the point at which the stress is being calculated (in meters, m),
- $I$ is the moment of inertia of the beam's cross-sectional area (in meters to the fourth power, m\(^4\)).

When selecting a beam it is important to be able to include a safety factor in the design. The safety factor is a multiplier that is applied to the calculated stress to ensure that the beam will not fail under the expected load.

If the stress on the beam is sigma, the required design specifications including the safety factor are given by:

$$ \sigma_{\text{yield}} = \text{Safety Factor} \times \sigma $$

Write Python code to do the following:

* Define a function called `beam_stress` which accepts three input arguments: `M`, `y`, and `I`, and a fourth optional argument called `safety_factor` which defaults to 1.0.
* Define a variable called `sigma` and set it equal to the result of the equation above.
* Define a variable called `sigma_yield` and set it equal to the result of the equation above.
* Have your function print the response: `The stress on the beam is: <sigma> N/m^2, and the required beam yield strength is <sigma_yield> N/m^2`, where sigma, and sigma yield are the integer values of the variable for the stress on the beam and yield strength respectively.
* Implement your function so that it returns the stress on a beam for the given inputs.

Provide an example usage of your function by calling it with the following inputs:
* `M = 1000` Nm
* `y = 0.1` m
* `I = 0.0001` m\(^4\)

Your code should print the response: 

Note: You can control the precision of the output by using string formatting to round the output to a specific number of decimal places. For example, to round the output to 2 decimal places, you can use the following code:

```python
print(f"{sigma:0.2f}")
```

Your code replaces the prompt:  `...`

```python
# Your function for beam_stress goes here
...

# Example usage:
M = ...
y = ...
I = ...
safety_factor = ...

# Call the function with example inputs save to a variable yield stress
...
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/adE3mdGL) add link to navigate to the assignment on GitHub classroom.

If you need further instructions on submitting your assignment please look at Lab 1. 

## Viewing your score

Each `.ipynb` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

We have both public and hidden tests. You will be able to see the score of both tests, but not the specific details of why the test passed or failed. 

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q0-Checking-Your-Name (points: None)
```
>>> def get_local_part(email):
...     return email.split('@')[0]
>>> local_part_email = get_local_part(drexel_email)
>>> assert local_part_email == drexel_id, f"The DrexelID extracted from the email '{local_part_email}' does not match the Drexel ID '{drexel_id}'."
>>> if local_part_email == drexel_id:
...     import drexel_jupyter_logger
...     assignment = 'Quiz6'
...     drexel_jupyter_logger.variable_logger_csv(assignment, 'info')
...     drexel_jupyter_logger.variable_logger_csv(first_name, 'info')
...     drexel_jupyter_logger.variable_logger_csv(last_name, 'info')
...     drexel_jupyter_logger.variable_logger_csv(drexel_id, 'info')
...     drexel_jupyter_logger.variable_logger_csv(drexel_email, 'info')
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_1')
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q2_1')
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q3_1')
...     drexel_jupyter_logger.variable_logger_csv('0, 2', 'q4_1')
...     drexel_jupyter_logger.variable_logger_csv('0, 2', 'q5_1')
...     drexel_jupyter_logger.variable_logger_csv('0, 2', 'q6_1')
...     drexel_jupyter_logger.variable_logger_csv('0, 3', 'q7_1')
```
failure_message: It does not appear that you entered your information correctly. Please try again.
success_message: You have entered your information correctly.


Tests for q1-stress-on-beam (points: 12)
```
>>> import drexel_jupyter_logger
>>> assert M == 1000, 'The value of M is not correct.'
>>> assert y == 0.1, 'The value of y is not correct.'
>>> assert I == 0.0001, 'The value of I is not correct.'
>>> assert safety_factor == 1.5, 'The value of the safety factor is not correct.'
>>> if M == 1000 and y == 0.1 and (I == 0.0001) and (safety_factor == 1.5):
...     drexel_jupyter_logger.variable_logger_csv('1, 1', 'q1_1')
```
success_message: Test values are implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> with patch('builtins.print') as mock_print:
...     assert yield_stress == 1500000, 'The stress is not calculated correctly.'
...     if yield_stress == 1500000:
...         drexel_jupyter_logger.variable_logger_csv('1, 1', 'q2_1')
```
success_message: Test for stress implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> with patch('builtins.print') as mock_print:
...     assert type(beam_stress(1, 2, 3, 1)) == float, 'The function does not return a numerical value of type float.'
...     if type(beam_stress(1, 2, 3, 1)) == float:
...         drexel_jupyter_logger.variable_logger_csv('1, 1', 'q3_1')
```
success_message: Function correctly returns a numerical value.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> (M, y, I, safety_factor) = (2000, 0.4, 0.002, 4)
>>> with patch('builtins.print') as mock_print:
...     assert beam_stress(M, y, I) == 400000.0, 'yeild stress with default safety factor is not correct.'
...     if beam_stress(M, y, I) == 400000.0:
...         drexel_jupyter_logger.variable_logger_csv('2, 2', 'q4_1')
```
success_message: default safety factor is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> (M, y, I, safety_factor) = (2000, 0.4, 0.002, 4)
>>> with patch('builtins.print') as mock_print:
...     out = beam_stress(M, y, I, safety_factor)
...     assert out == 1600000, 'Safety factor implemented correctly.'
...     if out == 1600000:
...         drexel_jupyter_logger.variable_logger_csv('2, 2', 'q5_1')
```
success_message: safety factor is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> import inspect
>>> signature = inspect.signature(beam_stress)
>>> parameters = signature.parameters
>>> assert len(parameters) == 4, f'Expected 4 parameters, found {len(parameters)}'
>>> safety_factor_param = parameters.get('safety_factor')
>>> assert safety_factor_param is not None, 'Safety factor parameter missing.'
>>> assert safety_factor_param.default != inspect.Parameter.empty, 'Safety factor does not have a default value, thus not optional.'
>>> if len(parameters) == 4 and safety_factor_param is not None and (safety_factor_param.default != inspect.Parameter.empty):
...     drexel_jupyter_logger.variable_logger_csv('2, 2', 'q6_1')
```
success_message: Input parameters is is correctly defined.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> expected_message = 'The stress on the beam is: 400000 N/m^2, and the required beam yield strength is 400000 N/m^2'
>>> with patch('builtins.print') as mock_print:
...     (M, y, I) = (2000, 0.4, 0.002)
...     out = beam_stress(M, y, I)
...     mock_print.assert_called_once_with(expected_message)
>>> if mock_print.call_args[0][0] == expected_message:
...     drexel_jupyter_logger.variable_logger_csv('3, 3', 'q7_1')
```
success_message: Input parameters is is correctly defined.
