# ⌛️ Quiz 3 - Using Python to Solve Math Problems

This quiz will explore using python's capabilities for mathematical calculations.


## Instructions

You will click on the following link to start your quiz. Do not start the quiz before your lab section. We have checks in place to ensure that you do not start the quiz early. If you do, you will receive a 0 on the quiz. Furthermore, it could be considered an academic integrity violation.

[CLICK HERE TO START YOUR QUIZ IN LAB](https://online.wiseattend.com/student/testLink?c=ENGR131&k=2801) 



## Part 1: Entering Your Information for Credit

To receive credit for assignments it is important we can identify your work from others. To do this we will ask you to enter your information in the following code block.

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
# Please provide your first name, last name, Drexel ID, and Drexel email. Make sure these are provided as strings. "STRINGS ARE TEXT ENCLOSED IN QUOTATION MARKS."

# In the assignments you will see sections of code that you need to fill in that are marked with ... (three dots). Replace the ... with your code.
first_name = ...
last_name = ...
drexel_id = ...
drexel_email = ...

```

**Question 1: Calculate the voltage across a capacitor** 

The charging or discharging behavior of a capacitor in a resistor-capacitor circuit is given by

$$ V_{\text{cap}}(t) = V_{\text{supply}} \cdot \left(1 - e^{-\frac{t}{RC}}\right)
 $$

where $V(t)$ is the voltage across the capacitor at time $t$ in seconds, $V_{\text{supply}}$ is the supply voltage in volts, $R$ is the resistance in ohms, and $C$ is the capacitance in farads.

Write python code to do the following:

* Define variables: `V_supply = 12`, `R = 1000`, and `C = 0.001`. 
* Calculate the voltage across the capacitor at `t = 5` seconds based on the equation above.
* Store the final value in a variable called `V_cap`.


Note: Your code should be indented under what's called the function definition: `def question_1():`. This helps the autograder run all the code you write by calling a single name. (Functions will be taught in future weeks.). 

Note: Python cares about **Tabs** make sure the spacing is correct. If you get an `Indention Error` it likely means your spacing is not correct. 

Your solution should look like this: 

```python
def question ():
    line1 = line1_code # line 1 of code
    line2 = line2_code # line 2 of code

    return x, y, z # variables to return separated by commas
```

Your code replaces the prompt:  `...`

```python
import numpy as np

# define a function called question_1 to be used for grading
def question_1():

    ...

    # The following line outputs these values from the function so that they can be accessed by the grader
    return t, V_supply, R, C, V_cap

```

**Question 2: Using Python as a calculator** 

We have seen how Python can be used to evaluate mathematical expressions. This problem provides practice incorporating various mathematical functions and constants.  


Write python code to do the following:

* Define variables: `a = 2`, `b = -7`, and `x = 3`
* For each of the quantities $E_0$ through $E_2$, construct a one-line Python expression that computes the value and assigns it to a variable. (You should perform your computations using the variable names $a$, $b$, and $c$ instead of their numerical values). In Python, we will use variable names `E0` through `E2` to store the values $E_0$ through $E_2$. Definitions of the variables `E0` through `E2` have been started for you in the template.

$\text{i.  } E_0 = \arcsin(\sin(a))$

$\text{ii.  } E_1 = \cos(b^2)$

$\text{iii.  } E_2 = \log_{10} \left(\left|3x\right|\right)$

You should use the `NumPy` library for the required mathematical functions. You can discover the necessary functions using the [Numpy Documentation](https://numpy.org/doc/stable/reference/)

We have used the standard `NumPy` convention `import numpy as np`, thus, for example the `cos` function is used by typing `np.cos(<value>)`

Again, your code should be indented under the function definition: `def question_3():`.

Your code replaces the prompt:  `...`

```python
# import the np library to be used
import numpy as np

# define a function called question_3 to be used for grading
def question_3():

    a = ...
    b = ...
    x = ...

    E0 = ...
    E1 = ...
    E2 = ...

    # output these values from the function so that they can be accessed by the grader
    return a, b, x, E0, E1, E2
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/Ogj3mdLq) to navigate to the assignment on GitHub classroom.

**Please submit just your output.log file. No other files are needed.**


If you need further instructions on submitting your assignment, please ask your TA during lab. 

## Viewing your score

Each `log` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

**Make sure you run your code, and run the tests. If you do not run the test you will not get credit for your work.**

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q0-Checking-Your-Name (points: None)
```
>>> def get_local_part(email):
...     return email.split('@')[0]
>>> local_part_email = get_local_part(drexel_email)
>>> assert local_part_email == drexel_id, f"The DrexelID extracted from the email '{local_part_email}' does not match the Drexel ID '{drexel_id}'."
>>> if local_part_email == drexel_id:
...     import drexel_jupyter_logger
...     drexel_jupyter_logger.variable_logger_csv(first_name, 'info')
...     drexel_jupyter_logger.variable_logger_csv(last_name, 'info')
...     drexel_jupyter_logger.variable_logger_csv(drexel_id, 'info')
...     drexel_jupyter_logger.variable_logger_csv(drexel_email, 'info')
```
failure_message: It does not appear that you entered your information correctly. Please try again.
success_message: You have entered your information correctly.


Tests for q1-interest (points: 2)
```
>>> import drexel_jupyter_logger
>>> (t, V_supply, R, C, V_cap) = question_1()
>>> assert t == 5 and V_supply == 12 and (R == 1000) and (C == 0.001)
>>> if t == 5 and V_supply == 12 and (R == 1000) and (C == 0.001):
...     drexel_jupyter_logger.variable_logger_csv('1, 1', 'q1_1')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_1')
```
failure_message: Variables incorrectly defined.
success_message: Variables correctly defined.
```
>>> import numpy as np
>>> import drexel_jupyter_logger
>>> (t, V_supply, R, C, V_cap) = question_1()
>>> assert np.isclose(sum([t, V_supply, R, C, V_cap]), 1028.920144636011)
>>> if np.isclose(sum([t, V_supply, R, C, V_cap]), 1028.920144636011):
...     drexel_jupyter_logger.variable_logger_csv('1, 1', 'q1_2')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q1_2')
```
failure_message: Final answer is incorrect.
success_message: Final answer is correct.


Tests for q3-calculator (points: 11)
```
>>> (a, b, x, E0, E1, E2) = question_3()
>>> import drexel_jupyter_logger
>>> assert a == 2 and b == -7 and (x == 3)
>>> if a == 2 and b == -7 and (x == 3):
...     drexel_jupyter_logger.variable_logger_csv('1, 1', 'q2_1')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q2_1')
```
failure_message: Variables incorrectly defined.
success_message: Variables correctly defined.
```
>>> import numpy as np
>>> import drexel_jupyter_logger
>>> (a, b, x, E0, E1, E2) = question_3()
>>> assert np.isclose(sum([a, b, x, E0]), -0.8584073464102067)
>>> if np.isclose(sum([a, b, x, E0]), -0.8584073464102067):
...     drexel_jupyter_logger.variable_logger_csv('1, 1', 'q2_2')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q2_2')
```
failure_message: E0 answer incorrect.
success_message: E0 answer correct.
```
>>> import drexel_jupyter_logger
>>> import numpy as np
>>> (a, b, x, E0, E1, E2) = question_3()
>>> assert np.isclose(sum([a, b, x, E1]), -1.6994074562563628)
>>> if np.isclose(sum([a, b, x, E1]), -1.6994074562563628):
...     drexel_jupyter_logger.variable_logger_csv('1, 1', 'q2_3')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q2_3')
```
failure_message: E1 answer incorrect.
success_message: E1 answer correct.
```
>>> import drexel_jupyter_logger
>>> import numpy as np
>>> (a, b, x, E0, E1, E2) = question_3()
>>> assert np.isclose(sum([a, b, x, E2]), -1.0457574905606752)
>>> if np.isclose(sum([a, b, x, E2]), -1.0457574905606752):
...     drexel_jupyter_logger.variable_logger_csv('1, 1', 'q2_4')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 1', 'q2_4')
```
failure_message: E2 answer incorrect.
success_message: E2 answer correct.
