```python
# Initialize Otter
import otter
grader = otter.Notebook("Final_Exam_Practice.ipynb")
```

# ⌛️ Practice Final Exam

This Practice Final Exam is designed to test your mastery of the Python programming language. You will be asked to respond to a series of questions both multiple choice and short answer problems using python. 10% of the points earned on this practice final prior to the exam will be added to your final exam score. 


## Instructions

You are allowed to use anything with the python kernel to help you solve the problems. You must run all code and press the submit button when shown to record your grades. If you do not do this your scores might not be recorded. All of your responses should be persistent if you restart the kernel, however entries can only be updated if you run the cell and hit submit. 


## Part 1: Entering Your Information for Credit

To receive credit for assignments, it is important we can identify your work from others. To do this, we will ask you to enter your information in the following code block.

You must hit submit and see the message saying the data is saved. If you do not see this message contact your instructor to ensure you are getting credit for your work.

**Note: do not delete the output.log file or the .responses.json file or you might lose your current work.**

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
import pkg_resources
from subprocess import call
import sys

package_name = 'ENGR131_Util_2024'
version = '1.0.0'
package_version = f'{package_name}=={version}'

try:
    # Check if the package and version are installed
    pkg_resources.require(package_version)
    print(f'{package_version} is already installed.')
except pkg_resources.DistributionNotFound:
    # If not installed, install the package
    print(f'{package_version} not found. Installing...')
    call([sys.executable, '-m', 'pip', 'install', package_version])
except pkg_resources.VersionConflict:
    # If a different version is installed, you can choose to upgrade/downgrade
    installed_packages = {dist.key: dist.version for dist in pkg_resources.working_set}
    installed_version = installed_packages.get(package_name.lower())
    print(f'{package_name} {installed_version} is installed, but {version} is required.')
    # Optionally, upgrade or downgrade the package to the required version
    call([sys.executable, '-m', 'pip', 'install', '--upgrade', package_version])
```

```python
from ENGR131_Util_2024 import responses, StudentInfoForm, cell_logger
from ENGR131_Util_2024 import submit_question, ResponseStore, ValidateLogFile


# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', cell_logger)
responses["assignment"] = "practicefinal_1"

StudentInfoForm(**responses)
```

**Question 1: Match the Correct Definition**

*18 points*

Please run the block of code below and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

```python
# from ENGR131_Util_2024 import DataTypes
from ENGR131_Util_2024 import responses
from ENGR131_Util_2024 import TypesQuestion

TypesQuestion(**responses)
```

**Question 2: Select the option that matches the definition** 

*9 points*

Please run the block of code below and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

```python
from ENGR131_Util_2024 import MCQQuestion
from ENGR131_Util_2024 import responses

MCQQuestion(**responses)
```

**Question 3: Select all statements which are TRUE** 

*20 points*

Please run the block of code below and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

```python
from ENGR131_Util_2024 import responses
from ENGR131_Util_2024 import SelectMany


SelectMany(**responses)
```

**Question 4: Reading and Interpreting Python Code** 

*50 points*

You are provided with the following code:

```python
1   numbers = [5, 2, 3.0]
2   even = 0
3   odd = 0
4
5   while odd < 6:
6       for num in numbers:
7           if num % 2 == 0:
8               even += num
9           else:
10              odd += num
```

Please run the block of code below and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

1. You will be asked to select the most relevant comment for the lines of code selected. 
2. You will be asked to describe how the code is executed. Please use the debugger tool to help ensure you have a correct answer.

   1. Line number: This is the line number of the code that is run during the execution of the code.
   2. Variable Changed: This is the variable whose value is changed during the execution of the code. *Note* if no variable is changed - such as when evaluating logical expressions, select "None".
   3. Current Value: This is the value of the variable after the line of code is executed. If a logical expression is evaluated, please enter the value of the logical expression.
   4. DataType: This is the data type of the variable after the line of code is executed. If a logical expression is evaluated, please enter the data type of the logical expression.
   
**Note**: The autograder has mechanisms to provide partial credit. If you skip a step, it will continue grading at the next correct step. 

We provide the code in Python for your convenience.

```python
numbers = [5, 2, 3.0]
even = 0
odd = 0

while odd < 6:
   for num in numbers:
      if num % 2 == 0:
         even += num
      else:
         odd += num
```

```python
from ENGR131_Util_2024 import ReadingPythonQuestion
from ENGR131_Util_2024 import responses

ReadingPythonQuestion(**responses)
```

## Question 5: Making a Graph 

*35 points*

A friend gives you a card with the following function written on it. You need to figure out what they are saying. Implement a function to reveal the hidden message.

1. Import the packages you need.

   - You will need `numpy` import this as the variable `np`
   - You will need the `pyplot` module from `matplotlib`. Import `pyplot` as a variable `plt`.


2. Make a function called `draw_card_eq`.
   - This should take one input `a`.
   - Create a linearly-spaced vector using the `linspace` method in `np`.
     - The vector should go from 0 to $2\pi$ with 2000 steps.
     - Assign this vector to the variable `t`.
   - Write the equation for `x`:
     $$ x = a(16(\sin(t))^3)$$
   - Write the equation for `y`:
     $$ y = a(13\cos(t) - 5\cos(2t) - 2\cos(3t) - \cos(4t))$$

   - Use the `plot` method in `plt` to plot `x` vs `y`, set the plot color to `red` using the input `red` as a string
     - Save this to the variable `plot_`.
   - Return `plot_`.

3. Call the function `draw_card_eq` with `a` equal to 1. Doing so will reveal the message. 
   - Save the output of the function to a variable `plot_`.


```python
# Your import statements go here.
# We would recommend checking this with the autograder.
# Completion of this step should pass tests 1-2.
...

# Your function goes here
# We would recommend checking this with the autograder.
# Completion of this step should pass test 3.
...

# Call your function here
# Make sure your check your code with the autograder. If you do not run the autograder your responses will not be graded.
# Completion of this step pass all tests.
...
```

```python
grader.check("q5 - Making a Graph")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results are aaved. 

```python
submit_question()
```

## Question 6: Using Data Types (30 Points)

We have written a function that converts a DataType. Pass values with the expected DataTypes as defined below to the `convert_to_string` function. The function will print the output. 

0. Run the `convert_to_string` function first.
1. Pass an integer `42` to `convert_to_string`. 
    - Save the output to a variable `int_`.
2. Pass an float `3.14159` to `convert_to_string`. 
    - Save the output to a variable `float_`.
3. Pass an boolean with a true value to `convert_to_string`. 
    - Save the output to a variable `bool_`.
4. Pass an string `I love coding` to `convert_to_string`. 
    - Save the output to a variable `string_`.
5. Pass a list of integers `1`, `2`, `3` to `convert_to_string`. 
    - Save the output to a variable `list_`.
6. Pass a tuple of values `a`, `b`, `c` to `convert_to_string`. 
    - Save the output to a variable `tuple_`.
7. Pass a dictionary with a key `name` and a value `jay` as a string and a key `age` and a value `12` as an integer to `convert_to_string`. 
    - Save the output to a variable `dict_`.



```python
def convert_to_string(data):
    if isinstance(data, int):
        print(data)
    elif isinstance(data, float):
        print("{:.2f}".format(data))
    elif isinstance(data, bool):
        print(str(data))
    elif isinstance(data, str):
        print(data)
    elif isinstance(data, list):
        elements = ", ".join([str(x) for x in data])
        print(elements)
    elif isinstance(data, tuple):
        elements = ", ".join([str(x) for x in data])
        print(elements)
    elif isinstance(data, dict):
        pairs = ", ".join(
            ["Key = " + str(k) + " Value = " + str(v) for k, v in data.items()]
        )
        print(pairs)
    return data
```

```python
# 1. Pass an integer `42` to `convert_to_string`. 
# Save the output to a variable `int_`.
# completion of this step will pass test 1
...
```

```python
# 2. Pass an float `3.14159` to `convert_to_string`. 
# Save the output to a variable `float_`.
# completion of this step will pass test 2
...
```

```python
# 3. Pass an boolean with a true value to `convert_to_string`. 
# Save the output to a variable `bool_`.
# completion of this step will pass test 3
...
```

```python
# 4. Pass an string `I love coding` to `convert_to_string`. 
# Save the output to a variable `string_`.
# completion of this step will pass test 4
...
```

```python
# 5. Pass a list of integers `1`, `2`, `3` to `convert_to_string`. 
# Save the output to a variable `list_`.
# completion of this step will pass test 5
...
```

```python
# 6. Pass a tuple of values `a`, `b`, `c` to `convert_to_string`. 
# Save the output to a variable `tuple_`.
# completion of this step will pass test 6
...
```

```python
# 7. Pass a dictionary with a key `name` and a value `jay` as a string and a key `age` and a value `12` as an integer to `convert_to_string`. 
# Save the output to a variable `dict_`.
# completion of this step will pass test 7
...
```

```python
grader.check("q6 - Exploring DataTypes")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results are aaved. 

```python
submit_question()
```

## Question 7: Building a Thermometer (36 Points)

In chemical engineering, it is common to monitor the temperature of a reactor. We will build a class to implement a temperature sensor. This sensor will save and also set a temperature value in Celsius. Additionally, the class will convert the temperature given in Celsius to Fahrenheit.

1. Build a class called `Thermometer`. This class will have to be initialized and have two methods.
    - Add an initialization method that accepts an initial temperature in Celsius as a variable `temperature`.
      - Save the `temperature` as an object attribute `temperature_c`. This is the temperature in $^{\circ}C$ .

    - Build a method `set_temperature` that will set the object attribute `temperature_c`.
      - This method takes one input `temperature`.
      - We want to make sure the input is a number. Do this by **trying** to convert it to a `float` using the built in `float` method.
        - If `temperature` is a number that can be converted to a float, save it to the object attribute `temperature_c`.
       - If `temperature` cannot be converted to a float, set the object attribute `temperature_c` to the string 'fail'.

    - Build a method `to_fahrenheit` that converts the temperature from Celsius to Fahrenheit.
      - Convert the temperature using the following equation.
        $$Fahrenheit = \frac{9}{5}Celsius +32 $$
      - Return the temperature in Fahrenheit.  

2. Instantiate a `Thermometer` `water_sensor` with an initial `temperature` of `12` $^{\circ}C$.
   
3. Extract the temperature in $^{\circ}C$ from the object `water_sensor`.
    - Assign the output to a variable `temp_c`.

4. Extract the temperature in $^{\circ}F$ from the object `water_sensor`.
    - Assign the output to a variable `temp_f`.


```python
# Your class for Thermometer goes here
...

# Instantiate the Thermometer water_sensor here
...

# extract the temperature in degrees celsius
...

# extract the temperature in degrees fahrenheit
...
```

```python
grader.check("q7-Building a Thermometer")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results are aaved. 

```python
submit_question()
```

## Question 8: Determining a Price for a Circuit (40 Points)

As an engineer, you might want to estimate the total price at which you would sell a circuit you designed based on the cost of goods from a supplier and a desired profit margin. We will write a Python function to assist you in this task.

We have provided you with two dictionaries, `parts_cost` and `circuit_parts`. `parts_cost` defines the the parts from the supplier, and `circuit_parts` defines the components needed to build the circuit. The circuit uses two `Custom Processor` components, which are not available from the supplier, that each costs $12.

*Cost of Goods from Supplier*

| Part Name (Key)     | Cost per Item (value, in $) |
| ---------------- | -------------- |
| 10-ohm resistor  | 0.025          |
| 100-ohm resistor | 0.03           |
| 1-pF capacitor   | 0.05           |
| 10-pF capacitor  | 0.12           |

*Components Required for Building the Circuit*

| Part Name  (Key)       | Quantity (sub-key) | Cost to Build (value, in $)|
| ---------------- | -------- | -------------- |
| 10-ohm resistor  | 10       | N/A             |
| 10-pF capacitor  | 3        | N/A            |
| 1-pF capacitor   | 7        | N/A            |
| Custom Processor | 2        | 12            |

1. Import the package `numpy` as `np`. 
    - You might not need to do this step if you have done it as part of the previous questions; however, it is good practice.
2. Define a function `price_calculator` that accepts three inputs: `parts_cost`, `circuit_parts`, and `profit`.
    - Initialize a variable `cost_of_goods`, and set it equal to 0.
    - Using a loop, compute the total `cost_of_goods`. 
        - You should use the `.items()` method of a dictionary that returns an iterator, where each iteration is a tuple containing the key and the value. 
        - You will need to use a branching statement to get the cost per item from the dictionary with the cost per item or the nested dictionary with the cost to build, whichever is appropriate for the part.
    - Determine the `sales_price` by adding the profit to the `cost_of_goods`. 
        - Multiply the `cost_of_goods` by (1 + `profit`). 
        - Using the `np.round` method, round the price to 2 decimal places. 
        - Save this value to the variable `sales_price`.
    - Call the `print_total` function within `price_calculator`.
        - Save the output to the variable `order_total`.
    - Print the variable `order_total` within the `price_calculator` function.
    - Return the variable `sales_price`.
3. Call the function `price_calculator` with the inputs `parts_cost`, `circuit_parts`, `profit`. 
    - Save the output to the variable `total_price`.


```python
# Your import statement goes here.
...

parts_cost = {
    "10-ohm resistor": 0.025,
    "100-ohm resistor": 0.03,
    "1-pF capacitor": 0.05,
    "10-pF capacitor": 0.12,
}

circuit_parts = {
    "10-ohm resistor": {"quantity": 10},
    "10-pF capacitor": {"quantity": 3},
    "1-pF capacitor": {"quantity": 7},
    "Custom Processor": {"quantity": 2, "price": 12},
}

# profit assignment goes here, we have set this to 63%
profit = 0.63

# We have provided you with a function to print the total cost of the order
# you need to call this function but do not need to change it.
def print_total(sales_price):
    return f"Total cost of the order: ${sales_price}"


# Your function price_calculator goes here
# Completion of this step should pass test 3
...

# Your call of the function price_calculator goes here
# Completion of this step should pass all tests
...
```

```python
grader.check("q8-Determining a Price for a Circuit")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results are aaved. 

```python
submit_question()
```

## Question 9: Plotting and Fitting (60 Points):

One of the most common tasks in engineering is fitting data to a model. For this question we will generate some noisy data with a gaussian distribution - one of the most common statistical distributions. We will then plot this data, and fit this data to the model.

1. Start by importing the necessary tools. If you imported any of these packages before, you might not need to do so again; however, it is good practice.

   - Import `numpy` as the variable `np`.
   - Import the `pyplot` module from `matplotlib` as a variable `plt`. 
   - Import the function `curve_fit` from `scipy.optimize`.

2. Define a function `gaussian` that computes the Gaussian function using the following equation:
   $$Ae^{\frac{-(x-b)^2}{2c^2}}$$
   - The function should accept inputs of `x`, `a`, `b`, and `c`.

3. Use the `linspace` method in `numpy` to generate a vector from `-20` to `20` with `100` steps. 
   - Save this vector to the variable `x_data`.

4. Use the defined `gaussian` function to compute the true, resulting value when `a` is `3`, `b` is `4`, and `c` is `7`. 
   - Save this to the variable `y_true`.

5. Generate noisy data using the `random.normal` method in `numpy`. 
   - The noisy data should be the same size as the `x_data`. 
   - Scale the noise by a factor of `0.3`, and add the noise the `y_true` variable.
   - Assign the output to `y_noisy`.

6. Using the `scatter` method from `matplotlib` `pyplot`, plot the `x_data` vs `y_noisy`. 
   - Add a label, "Noisy Data". 
   - Save this plot to the variable `plot_1`.

7. Use the imported `curve_fit` method to fit the noisy data to the `gaussian` function over the range of `x_data`. 
   - Save the fit results to the variable `popt`, and the parameter covariances to a variable `pcov`. 
     - **Do not provide any initial guess**, or provide `None`. If you provide an initial guess, the autograder might fail.

8. Calculate the fitted curve using the defined `gaussian` function. 
   - You can use the \* method to unpack the list of fitted parameters if you like. 
   - Save the result to a variable `y_fit`.

9.  Plot the fit results as `x_data` vs `y_fit` using the `plot` method in `pyplot`. 
    - Set the color of the line to green using the `g` tag.
    - Add a label, "Fitted Gaussian". 
    - Save this plot to the variable `plot_2`.

10. Using the methods within `plt`, set the `x_label` and `y_label` to `x` and `y`, respectively.
    
11. Use the `legend` method in `plt` to add a legend.


```python
# 1. Add your import statements here. 
# Completion of this step will pass the first test
...

# do not delete the line of code below, it will cause problems with the autograder.
np.random.seed(42)

# 2. Define the Gaussian function.
# completion of this step will pass test 2.
...

# 3. define your x_data vector
# completion of this step will pass test 3.
...

# 4. Call the gaussian function as described above
# save the output to the variable y_true
# completion of this step will pass test 4.
...

#5. Generate the noisy data described above
# save the output to the variable y_noisy
# completion of this step will pass test 5.
...

# 6. Plot the noisy data using a scatterplot
# make sure to add the label "Noisy Data"
# save the output to the variable plot_1
...

# 7. Fit the noisy data to a Gaussian function using curve_fit
# save the outputs to the variables popt and pcov
# completion of this step will pass test 6.
...

# 8. Calculate the fitted curve using the gaussian function
# save the output to the variable y_fit
# completion of this step will pass test 7. 
...

# 9. Plot the fitted results as described in the instructions
# save the output to the variable plot_2
# completion of this step will pass several more tests. 
# make sure to add the label "Fitted Gaussian"
...

# 10. Add axis labels to the plot
# completion of this step will pass a test.
...

# 11. Add a legend to the plot
# completion of this step will pass all tests.
...
```

```python
grader.check("q9-Plotting and Fitting")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results are aaved. 

```python
submit_question()
```

## Validating your log file

Before your submit your results we would recommend running the following command to validate your log file. This will ensure that you have followed the correct format and that your log file is valid. 

```python
ValidateLogFile("./output.log", responses["assignment"], 4)
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/ClDJl_RS) to navigate to the assignment on GitHub classroom.

**Please submit just your output.log file. No other files are needed.**


If you need further instructions on submitting your assignment, please ask your TA during lab. 

## Viewing your score

Each `log` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

**Make sure you run your code, and run the tests. If you do not run the test you will not get credit for your work.**

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q5 - Making a Graph (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [3, 3, 3, 6, 14, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q5_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_1'
>>> max_score = 3
>>> score = 0
>>> if np.__version__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': np.__version__}
>>> responder.add_response(response)
>>> assert np.__version__ is not None, 'numpy is incorrectly imported'
```
failure_message: numpy incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_2'
>>> max_score = 3
>>> score = 0
>>> if plt.__package__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': plt.__package__}
>>> responder.add_response(response)
>>> assert plt.__package__ is not None, 'plt is incorrectly imported'
```
failure_message: plt random incorrectly implemented.
success_message: plt correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_3'
>>> max_score = 3
>>> score = 0
>>> if callable(draw_card_eq):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': callable(draw_card_eq)}
>>> responder.add_response(response)
>>> assert callable(draw_card_eq), 'draw_card_eq is not a function'
```
failure_message: draw_card_eq is not a function
success_message: draw_card_eq is callable.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import matplotlib.pyplot as plts
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_4'
>>> max_score = 6
>>> score = 0
>>> if isinstance(plot_[0], plts.Line2D):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': isinstance(plot_[0], plts.Line2D)}
>>> scorer.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     scorer.submit()
>>> scorer = submit_score()
>>> assert isinstance(plot_[0], plts.Line2D), 'You did not return a plot object.'
```
failure_message: You did not return a plot object.
success_message: You successfully returned a plot object.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_5'
>>> max_score = 14
>>> score = 0
>>> if np.isclose(sum(plot_[0].get_xdata()[::3]), 1.6561871562318847e-07):
...     score += int(max_score / 2)
>>> if np.isclose(sum(plot_[0].get_ydata()[::4]), 1.2500506348082654):
...     score += int(max_score / 2)
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': sum(plot_[0].get_xdata()[::3])}
>>> responder.add_response(response)
>>> assert np.isclose(sum(plot_[0].get_xdata()[::3]), 1.6561871562318847e-07), 'The values for x_data in the plot are incorrect. '
>>> assert np.isclose(sum(plot_[0].get_ydata()[::4]), 1.2500506348082654), 'The values for y_data in the plot are incorrect. '
```
failure_message: The values for your plot are incorrect.
success_message: The values for your plot are correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q5_6'
>>> max_score = 6
>>> score = 0
>>> if plot_[0].get_color() == 'red' or plot_[0].get_color() == (1.0, 0.0, 0.0, 1.0) or plot_[0].get_color() == 'r':
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': plot_[0].get_color()}
>>> responder.add_response(response)
>>> assert plot_[0].get_color() == 'red' or plot_[0].get_color() == (1.0, 0.0, 0.0, 1.0) or plot_[0].get_color() == 'r', 'The line on the plot is not the correct color'
```
failure_message: The plot is not red.
success_message: The plot is red as expected.


Tests for q6 - Exploring DataTypes (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [3, 3, 3, 3, 6, 6, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q6_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_1'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDQ71dNmKvaPlZOj3cLT2FVfh66b09i0yV2Fl08bcFY0K_Io7ju7utXkX3cCmsb9HMsHTAHERX0X48acRp1MGHYXxSg==')
>>> if str(int_) == out and isinstance(int_, int):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': int_}
>>> responder.add_response(response)
>>> assert isinstance(int_, int)
>>> assert str(int_) == out, 'int_ incorrectly implemented.'
```
failure_message: int_ incorrectly implemented.
success_message: int_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_2'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRABl3uUfW0G0yLb5kCpVpz5XBRYUjowgJ9_c9CNAGielysMhb5cp72zDF0Gi9abZrATkUoZfUGv7UJp65Cm7WvjiQ==')
>>> if str(float_) == out and isinstance(float_, float):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': float_}
>>> responder.add_response(response)
>>> assert isinstance(float_, float)
>>> assert str(float_) == out, 'float_ incorrectly implemented.'
```
failure_message: float_ incorrectly implemented.
success_message: float_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_3'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRDR42csyMOkhel0VkaN7JXse5_3OFTam1ryIyiJvEdIcfDCfocojIiGZU2PPJQm-rrfBNXLyzMQ-QJgmZi7qEdthA==')
>>> if str(bool_) == out and isinstance(bool_, bool):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': bool_}
>>> responder.add_response(response)
>>> assert isinstance(bool_, bool)
>>> assert str(bool_) == out, 'bool_ incorrectly implemented.'
```
failure_message: bool_ incorrectly implemented.
success_message: bool_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_4'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRECpI3kPdhirMzwLAddFoJKDXs7mXJ5AWPYJTx_gTKGYY9VR0REjaPKLCAc0MQ5UuSr4cMcBBc8zSVg2GT6LiDvLg==')
>>> if str(string_) == out and isinstance(string_, str):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(string_)}
>>> responder.add_response(response)
>>> assert isinstance(string_, str)
>>> assert str(string_) == out, 'string_ incorrectly implemented.'
```
failure_message: string_ incorrectly implemented.
success_message: string_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_5'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRFLoBpFP_06PNjYcIdyOpLGFYYboKIjMjTAu936eLjoKByeaF5XTyYcvGh5PWyVVU1kFUqWj4YeqbRDfMTRpPYXGw==')
>>> if str(list_) == out and isinstance(list_, list):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(list_)}
>>> responder.add_response(response)
>>> assert isinstance(list_, list)
>>> assert str(list_) == out, 'list_ incorrectly implemented.'
```
failure_message: list_ incorrectly implemented.
success_message: list_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_6'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRGEX2Cl-qHsEg1NFnwyZSHVLsjlc1JJWn4Lj76LT7dsEoNqlBpXQry0zKyFS2oCSzdJbYSm95SKxB96pM_HHsyuNw==')
>>> if str(tuple_) == out and isinstance(tuple_, tuple):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(tuple_)}
>>> responder.add_response(response)
>>> assert isinstance(tuple_, tuple)
>>> assert str(tuple_) == out, 'tuple_ incorrectly implemented.'
```
failure_message: tuple_ incorrectly implemented.
success_message: tuple_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q6_7'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABkDRG1TzclExNe_t7qeJgaHYC9uWufJOfW5u7L198U2jtPxlFPQTnL9qj0MFHrNC7pIUfpz7od9FbuyQy6nuNnEFvlk2ZWGw3xRiU_wh4fVtRuV-Y=')
>>> if str(dict_) == out and isinstance(dict_, dict):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(dict_)}
>>> responder.add_response(response)
>>> assert isinstance(dict_, dict)
>>> assert str(dict_) == out, 'dict_ incorrectly implemented.'
```
failure_message: dict_ incorrectly implemented.
success_message: dict_ correctly implemented.


Tests for q7-Building a Thermometer (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [6, 6, 8, 8, 8]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q7_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_1'
>>> max_score = 6
>>> score = 0
>>> if isinstance(Thermometer, type):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(isinstance(Thermometer, type))}
>>> responder.add_response(response)
>>> assert isinstance(Thermometer, type), 'The Class Thermometer is not implemented.'
```
failure_message: The Class Thermometer is not implemented.
success_message: The Class Thermometer is implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_2'
>>> max_score = 6
>>> score = 0
>>> t = Thermometer(25)
>>> if hasattr(t, 'temperature_c') and t.temperature_c == 25:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(t.temperature_c)}
>>> responder.add_response(response)
>>> assert hasattr(t, 'temperature_c'), "Thermometer does not have 'temperature' attribute"
>>> assert t.temperature_c == 25, 'Thermometer initialization did not set attribute temperature correctly'
```
failure_message: The Class Thermometer initialization is implemented incorrectly.
success_message: The Class Thermometer initialization is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_3'
>>> max_score = 8
>>> score = 0
>>> t = Thermometer(25)
>>> t.set_temperature(50)
>>> if t.temperature_c == 50.0:
...     score += 4
>>> t_ = Thermometer(25)
>>> t_.set_temperature('abc')
>>> out = str(t_.temperature_c)
>>> if out == 'fail':
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([t.temperature_c, t_.temperature_c])}
>>> responder.add_response(response)
>>> t = Thermometer(25)
>>> t.set_temperature(50)
>>> assert t.temperature_c == 50.0, 'set temperature does not function properly'
>>> t = Thermometer(25)
>>> t.set_temperature('abc')
>>> out = str(t.temperature_c)
>>> assert out == 'fail', 'set temperature does not return the correct response when a unsuitable input is provided'
```
failure_message: set_temperature can not correctly handle non numeric inputs.
success_message: set_temperature can correctly handle non numeric inputs.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_4'
>>> max_score = 8
>>> score = 0
>>> t = Thermometer(0)
>>> if t.to_fahrenheit() == 32.0:
...     score += 4
>>> hold = t.to_fahrenheit()
>>> t.set_temperature(20)
>>> if t.to_fahrenheit() == 68.0:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([hold, t.to_fahrenheit()])}
>>> responder.add_response(response)
>>> t = Thermometer(0)
>>> assert t.to_fahrenheit() == 32.0, 'Incorrect conversion to Fahrenheit.'
>>> t.set_temperature(20)
>>> assert t.to_fahrenheit() == 68.0, 'Incorrect conversion to Fahrenheit.'
```
failure_message: Incorrect conversion to Fahrenheit.
success_message: Correct conversion to Fahrenheit.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q7_5'
>>> max_score = 8
>>> score = 0
>>> conditions = [hasattr(water_sensor, 'temperature_c'), hasattr(water_sensor, 'to_fahrenheit'), temp_c == 12, temp_f == 53.6]
>>> statements = ["Thermometer does not have 'temperature' attribute", "Thermometer does not have 'to_fahrenheit' attribute", 'Incorrect temperature for temp_c.', 'Incorrect temperature for temp_f.']
>>> for condition in conditions:
...     if condition:
...         score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: Incorrect conversion to Fahrenheit.
success_message: Correct conversion to Fahrenheit.


Tests for q8-Determining a Price for a Circuit (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [0, 1, 0, 6, 15, 12, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q8_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_1'
>>> max_score = 0
>>> score = 0
>>> conditions = [parts_cost['10-ohm resistor'] == 0.025, parts_cost['100-ohm resistor'] == 0.03, parts_cost['1-pF capacitor'] == 0.05, parts_cost['10-pF capacitor'] == 0.12]
>>> statements = ['10-ohm resistor not correctly defined', '100-ohm resistor not correctly defined', '1-pF capacitor not correctly defined', '10-pF capacitor not correctly defined']
>>> for condition in conditions:
...     if condition:
...         score += 0
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: parts_cost is incorrect. Please make sure you did not modify the dictionary.
success_message: parts_cost is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_2'
>>> max_score = 1
>>> score = 0
>>> if np.__version__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(np.__version__)}
>>> responder.add_response(response)
>>> assert np.__version__ is not None, 'numpy is incorrectly imported'
```
failure_message: numpy incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_3'
>>> max_score = 0
>>> score = 0
>>> conditions = [circuit_parts['10-ohm resistor']['quantity'] == 10, circuit_parts['10-pF capacitor']['quantity'] == 3, circuit_parts['1-pF capacitor']['quantity'] == 7, circuit_parts['Custom Processor']['quantity'] == 2, circuit_parts['Custom Processor']['price'] == 12]
>>> statements = ['10-ohm resistor quantity not correctly defined', '10-pF capacitor quantity not correctly defined', '1-pF capacitor quantity not correctly defined', 'Custom Processor quantity not correctly defined', 'Custom Processor price not correctly defined']
>>> for condition in conditions:
...     if condition:
...         score += 0
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: circuit_parts is incorrect. Please make sure you did not modify the dictionary.
success_message: circuit_parts is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_4'
>>> max_score = 6
>>> score = 0
>>> conditions = [price_calculator.__name__ == 'price_calculator', price_calculator.__code__.co_argcount == 3]
>>> statements = ['you did not name price_calculator correctly', 'price calculator does not have the correct number of inputs']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([price_calculator.__name__, price_calculator.__code__.co_argcount])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: price_calculator call is incorrect.
success_message: price_calculator call is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_5'
>>> max_score = 15
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkDz6Y_aio_akqMGkihfI5zKR2KkUg0yXVVHNrVV82qunlAwRjVZN2R2OYzUqkBBLBG-qd4t_2Vpx-iHxYPLMGTYQsFA==')
...     if str(price_calculator(parts_cost, circuit_parts, 0.2)) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(out)}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert str(price_calculator(parts_cost, circuit_parts, 0.2)) == out, 'price_calculator  computation is incorrect.'
```
failure_message: price_calculator computation is incorrect.
success_message: price_calculator computation is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_6'
>>> max_score = 12
>>> score = 0
>>> parts_2 = parts_cost.copy()
>>> parts_2['10-ohm resistor'] = 10
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD3dtfMyEiX9AYR1pum4OZ5xSbbZ367hae3wxg5SSAcor4m5CXO6u1OthJkNgyGQUotRmCY2v18f4rjxK7Z3kbCNC3A==')
...     if out == str(price_calculator(parts_2, circuit_parts, 0.2)):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(out)}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert out == str(price_calculator(parts_2, circuit_parts, 0.2)), 'Your code does not use the parts_cost correctly'
```
failure_message: Your code does not use the parts_cost correctly.
success_message: Your code correctly uses the parts_cost.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q8_7'
>>> max_score = 6
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD3D3Z-O5XJGgN67FRSetS5oiOetUTzWxQFbw_OTKnrjBEhnelTKcH0d22z4AneFL8pYrtOgSg0Q2qYc0A4dI-VdzzcV6pXrv5aKw-FAVmhHMZ7U=')
...     if print_total(price_calculator(parts_cost, circuit_parts, profit)) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(out)}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert print_total(price_calculator(parts_cost, circuit_parts, profit)) == out, 'The string being printed is not correct'
```
failure_message: print message is incorrect.
success_message: print message is correct.


Tests for q9-Plotting and Fitting (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> responder.create_file()
>>> points_ = [3, 9, 4, 6, 3, 9, 3, 4, 4, 3, 3, 6, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q9_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_1'
>>> max_score = 3
>>> score = 0
>>> conditions = [np.__name__ == 'numpy', plt.__name__ == 'matplotlib.pyplot', 'curve_fit' in dir()]
>>> statements = ['numpy not imported correctly', 'matplotlib.pyplot not imported correctly', 'scipy.optimize.curve_fit not imported correctly']
>>> for condition in conditions:
...     if condition:
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(plt.__name__)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: Your input statements are incorrect.
success_message: Your input statements are correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_2'
>>> max_score = 9
>>> score = 0
>>> x_test = np.array([0, 1, 2])
>>> (a_test, b_test, c_test) = (1, 1, 1)
>>> expected_output = np.array([0.60653066, 1.0, 0.60653066])
>>> conditions = [np.allclose(gaussian(x_test, a_test, b_test, c_test), expected_output)]
>>> statements = ['gaussian function is not implemented correctly']
>>> for condition in conditions:
...     if condition:
...         score += 9
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([*gaussian(x_test, a_test, b_test, c_test)])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: gaussian function is not implemented correctly
success_message: gaussian function is implemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_3'
>>> max_score = 4
>>> score = 0
>>> conditions = [np.allclose(x_data[0], -20), np.allclose(x_data[-1], 20), len(x_data) == 100, np.isclose(sum(x_data[3::3]), 19.99999999999997)]
>>> statements = ['x_data is not defined correctly', 'x_data is not defined correctly', 'x_data length is not correct', 'x_data is not defined correctly']
>>> for condition in conditions:
...     if condition:
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: x_data is not defined correctly.
success_message: x_data is defined correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_4'
>>> max_score = 6
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD-NzYrMglFaqNyS-zlkoQo-gwxc1YCD39uP-uNVNborNwKBB7YKbPEpayy9rq7hUGxJhZWrWmXBDQ14fxazkpMi5v-8ArClYcJqu2ZJvUs2LKzPwphZZ63pfZBxs8IK0DIw-EnU5mRoiaMT3YoTAD7rSGUrnDdqM44YGV3b_Gyt7mYE0n_obgjbXngI0DbuJzExiXeA4KQNpSsKJD6x870owlz7wNXADMhSrU03suFLlUx1pgGwJsg-7lQhUocs64ZahOJvSfABgqu_Ml_SDkaudvTyHTFCPy-bnq3xn-Cc4d9jmsdOxUmqMvl3xLZw1ALt6YrMBdlNZh0hBgjrT4iVP9tySJBG8dnov7TLz9_PuWRGRulKG-tYZU9XJ8uuFC09mXpDLE49gdlCDp_jYbvYfWX6rPlgBd7pXyQ2YWAKajlRzuJ1Q7TrXSVm_Xbj2eF48_iehGgz2cuzg0YARRKXvrg3JFdRn-jRhN9Pop4SoJVEcIbh6p6AmU0u2rUTFDN5UazGdLSWav8g2xQyHCm8_2i9il6lVF3_H4E5gdycNi6eb4hUV0i7XY1rUzfulJYmJzWzQk3KgbR54In9AfLWvgHY0IcBduStXBmpMfwIG7NMbbSPhElQ0D6gXjA3F7zpkahXZ3S0ihF0bfHvY32pIfRVORtMT811bYif8OPTJH2i3YcHiaS7ElIUJXkb4dnJTOOVCIB08EV0ONZgZHwtpsFCCwCFvRTouorcfz1SXs9a0vtjtau3CgBI3cl0yAGGuiazfipZAEsbB68QedFffR_QRl0ZBN_IfGIMC8ksxUw3R9J46tUET_doctLBvd7LFbDTfkECNoKCh-E5x-POLyjeWgscHe64WSRNkYgKaKTWP70Y9OdYg2CpCZseE55ovOqgNjTVCMlRhuQWO9zuBSLAU3NOfL-ms4n9AWrx75Wsg9gxkIGVm36HZ-yKO_wJDcrRRQqtYaS0noeqEdo3-Gmj4ldMgPOATkNXKLoFJ7P6FfnNYUiY7mPykn1_1tqlBEhVvSXrAteRNr6E_8s6lgFr4l2kDU8s8_XSgKu_wazGqw45WVMiO5Zgsh2SXh4VnnUE0LFubdThZsKCplGr_uJK2IFc0jiWlSid8rE2GiLU69Y0mVkwod2PQOfIfJRzeNkAA9raRbs5CkFO8XFVgLGjKWGIJw_8zKncKHZ-LgK0PVFJ2zMBMuN8fXofonS3v7Zm85gj16MGtfwIOVhvODG4BrOkZvI43TCM5pPRq1-5nqyJfJ6gpC0SWWwMrEF7YSZ64YIA5YmVlaquaB5Fam4p5O3EN8m83aQ4AcCBcRZAaimCtPIrIhK7SS394wkfQLFb8q_WUJe5RBHm3PoVsbjac3Vlu7rQTVByxSf9KdFNdtKTXJsMzLL0743ZMBmm4RA5UlQWE15AFeIyiSXP8B0TgFKe9J5g8W-eT8KWyqRw-x00_7lF0e-4tcB1B6CLOuk4kpkZk0WkucN0S1IzZRZsjXcOwmsoNDqFdpx7NrO08ZgmXqkSqwLCo99q0C6gpk2oX3-BBw9cDOgvFIpYE4A==')
...     if str(y_true) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(y_true)}
>>> responder.add_response(response)
>>> assert str(y_true) == out, 'y_true is incorrect'
```
failure_message: y_true is incorrect.
success_message: y_true is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_5'
>>> max_score = 3
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD-UVvxc6wECqC_U9taa6Art0EtaZg2i_WB0sp4FHCf3Ggyed8uUrljAkCPxg-AwYxN2jnqaPVgMh57F1OX1ftyjzR1SfhzBk5UAkxPaLIFHteLZY_7oK67GesWvX-jys-fhfKLktB_4ba54nXaVQP3K2xb60izc-42Y4VvddXu9Iggo5dy7VcJP8-6Ix_un8KqJo2J4CKYXFni2El-O9zrvvc1EcNYgvGdpzxFFeUHEHAP0AEqTp-Ow5D7T0x2iGcEPXtDUcmbgMG_t-HAQRr8ky9DNyuKx9k13pWT85O3jpLAtExhBYk2yXtzS9iU4pQEYGpYhdB2HQosUQcV5rXdgsC6pGcWlWhTab2uG32_Lt8-Y0jlb4VgALBAtxJGXuevGpJ7Pt1fzt7pjHikCVHcMbmkhLaWGzZq1uf6d6mbA-lxo1wGif7-ryFB4TTyze2HoxcGGgLI7vtuHgduBSd397OTTuS9PMMnR7oIXBZprj_NJrSmI79N4mMDlkWz-baZikhLw6LbThxinvd4HX6aQPVm3bPlV9M5Yc3lSqr2YQimUKqONnZGKxzcQx4TxfeRjdv8mS49Qey1KRU5zUlBS11NX-ScrqIFDAwFBsq6OpMSqd80rZquYE5E-IsHr0g4McffuDal20JZqXtP_kYMaT3Iv_KCa__9TKmpMHKJpQ44qhvzuBkqA59npHU05s5gbr5OpPwug1W_H7WTgjRYD-IaDU9cxxKWs9erfEkcCMh7peMvnBVwu2IY_W1u7ocNxXiMR2BJX6FbGO2P12t7D8vz3lyCJFHhkYggO06v9UF0L3-eX186gXNdegYnML3kVevZbOkRHFxvU2hoZ6dzLW8Z8SgaO9H1-Pac5Ima1Kh36EF1GH-93QGnSf6w2feXG9yByfSSbRharfHT6-79nYgl3EgGL0wo8imUaN8uwRhkxma8myGCAsNbkPKPiltTXfyv8yKEjc-qC5UivinYtFYYkhiDqIQvQDB4YURk0p4MxesM1rOgq6aML_jc7haRUcUes4ynGERvVt_Nc3bV03--BIzTZx4D2q3OMEUt25CyTU9oUr6r4C_JsrxhtK2d5U9qrLY6b-W-F-lBI3ZrRh_HcwHR9Hx7YTA7vRcXaOyYB-fUWjwXa2Gbxj1s9iEpzWf8lEisQk40FUhc2jSmtiZEn6MO3dO_YxJ3s9TE1Tmq9aK7gmb3zMbZn2gQyuHxzuvnUeJBi-DelNLvQXP7C3GYyMjA6TuajEzNY6VSVE7sKgQyw_vjE0QteWtkUpODlAbpT42M6laN3ahNeT7l2WE0D-wvjEHzflGOF7udat6QccodlHUge4DAJSS4vaTWs5J-aoSCmPNrYN65W3CIokMhxXS0VZ0n26ObdU384YVKdOTtUn796D7Mdk0ll0ImLLs_1RcihuZPkmEhE2rnWUINRBOLAyoQz6RrNGe-tC-uA3XNsxYbDQ9ReBPlU_R_BbkkgEzKtMV7juGtddtKDF-1lzhoahFX33CNMdWcuq189G87JsEED_tdmSloByFx3sEExB0uuj0--pnbH-CV1674xJ8x2WarCevNkOts1Aa0tDyE9v9v8NmX6cIKHsT8rkIhStCE1_Xknn2uMGlqNPp_5k30P5_mKeth6tti6uH-MMNO_qQ8V82Ki7X5u3HOmPgAa3_uJGD-JVvREQ8uNw8SVFpCmO-X207mkqYEFXO_Q=')
...     if str(y_noisy) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(len(y_noisy))}
>>> responder.add_response(response)
>>> assert str(y_noisy) == out, 'y_noisy is incorrect'
```
failure_message: y_noisy is not implmemented correctly
success_message: y_noisy is implmemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_6'
>>> max_score = 9
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD-VjEnvkOx1q5hIc6OiMyPE7iayGWAtZR_Kry0mvrQbCr0ffV9Lrck5j673YFnlWxRjo9qR44gKpYzUOzGhmCBpXKZ2Dec1hHWXhY4szHZZuirxAWmfMc66dj5dpA1bzjYFG')
...     chars_to_remove = ['[', ']', '\n']
...     for char in chars_to_remove:
...         out = out.replace(char, '')
...     out = [float(x) for x in out.split(' ')]
...     if np.allclose(popt, out):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([*popt])}
>>> responder.add_response(response)
>>> assert np.allclose(popt, out), 'popt is not returning the correct values'
```
failure_message: popt is not returning the correct values
success_message: popt is returning the correct values
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_7'
>>> max_score = 3
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABkD-WzDT72FrJF_4PrNsTIhnppKucv8pAjRnotXNrslY7YlSdDyqgiXabo32_bwcFXYkokXar7Vl-Bl1cN4lD70-PDCbfDPsxVDcGs25T4Z8IV-i6P7SAwaZ-k00NmVsJ4sF62ClW9ssomSd0jgYtbgy4J5sF9_KO-6s0ke6_Pfi1mrx8c90J5ZYsAwB9NKm45dZ-IWfjH955EkoxoPQcdqBQ16RG1J0k5DqGaP6hpRpR4mRANk50AFcgkX9PLqnKr3Z3AH1s_ugMEvza1GRlj1aVmaeF0JcHsVw1MxpZLLhPK0ikm8TcyX7cy6O8wi3xvjk5By8SFT96a0pqK5ObXGIuFAN6XlBH56yic8ttcRBODzaRiH_97wGvi8eUP0NNGOlMZFb_mPc7E9URWZCnJ0Krr57iOj5ALWDPqtMysnRGg4VCI9l6ZAFv0dP_HqbZT7U7TK_dFVAWNWv5MPeyg8rcGInJ2ygH1o1G8jRJRSbEcHQj3-ZU2K7DCFqc6xg2EO_31Dz6VLWnEBLVW2OUIth16vPCcjbmJgGuW0mgRulbogMNLFEXpGKJwi8LimBL_gZ0pee4vd8nZDKaRR8pSoLXsHyTHf-O8vXHzji4RiUF7VSReSrbpai1xoC7uyztjO1Vjce1bMf1pP7e3JNIzpJ7sfAlP0PPiERdG58Xe_2Qx_oWFMtQeDt-ViS4AWK76KlWgsiX0cFqWs3Dj_KMs1hmtywbEZtLkUjEUH6rclgw482Tp44fFQcUqBOOD92DJkGjmEr-FcVJtlu8R71hafSMBnOJ_cEUvZL-P5ljNUoUoM1I7REUK9hs-hiSYEyXbyCfv5dRSoXHOdNv0nDKoEcVWPYZPP3E74Qafmm47NHUNbcG3pMGXXMsOea7l5SnL_gfnYjNcvUIDEzZDXen2moiSP0P5EIeFTSItTdjSnM-TAMmDZlFSF3XtXD8vagpPb_-QdfPNOkkx1MQNjVXcKqDVPUwyuDkrqm0QFp2XrRW8Lwb7v5mENaD3HfZMlnpuBYAv9sXDtkb7Km30U9NgYExvScKEnBfbd61peROkRYXJiTA8uvx3xgWAeu47lCYNJ0jKomOQlJOVXH86TnPEjMKDQ9UUcFTJGcuQLaNfSPdSPqkf1gabg3e6VW4b--PSkhpss9OboUp1UOEo0KFnlj4Tl-zu8rwuWEUwlBUDYIyuKwWBYD_DpC62kBfCcuMuaM1xvbbMoedQLTm0OA0PIR_dP5nrW4IkGtnpbz43rqBMFCcGZ2un8KbCRT1mh8dB5nNdrtxl3dtBEB_nsl-M8gIqCruIfNHkoPx77SkXHT3zNnI66oYqgBpMCrTjShWB-xNtlRrKHqL0q9RpgPF9SJphcLBQ07eqV4s7TrJtfrhw41WuzLaCR7KYb_vsMj78KNrZmuaTDgGGu4lediwCYGOFA3I-jIxSj2jwYfDcrGKW9QUbDK5RTZ0_rEt6SwA5r7NpadXtv-RsGBycn-_vTNPQ_5PGIYGyY3GkQgVjiq0nBX1xI39PqmMKc0eyGKAH6pu2uhCw_FAe3GSSHzMke-zzIA==')
...     chars_to_remove = ['[', ']', ' \n']
...     for char in chars_to_remove:
...         out = out.replace(char, '')
...     out = out.replace('  ', ' ')
...     out = [float(x) for x in out.split(' ')]
...     if np.allclose(y_fit, out):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([*popt])}
>>> responder.add_response(response)
>>> assert np.allclose(y_fit, out), 'y_fit is returning incorrect values'
```
failure_message: y_fit is returning incorrect values
success_message: y_fit is returning correct values
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_8'
>>> max_score = 4
>>> score = 0
>>> conditions = [plot_1.axes.get_xlabel() == 'x', plot_1.axes.get_ylabel() == 'y']
>>> statements = ['xlabel is incorrect', 'ylabel is incorrect']
>>> for condition in conditions:
...     if condition:
...         score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([plot_1.axes.get_xlabel(), plot_1.axes.get_ylabel()])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: The plot labels are incorrect.
success_message: The plot labels are correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_9'
>>> max_score = 4
>>> score = 0
>>> conditions = [plot_1.get_label() == 'Noisy Data', plot_2[0].get_label() == 'Fitted Gaussian']
>>> statements = ['Noisy Data', 'plot_1 is not labeled correctly', 'Fitted Gaussian', 'plot_2 is not labeled correctly']
>>> for condition in conditions:
...     if condition:
...         score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([plot_1.get_label(), plot_2[0].get_label()])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: the plots are not labeled correctly
success_message: the plots are labeled correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_10'
>>> max_score = 3
>>> score = 0
>>> conditions = [plot_2[0].get_color() == (0.0, 0.5, 0.0, 1) or plot_2[0].get_color() == 'g' or plot_2[0].get_color() == 'green']
>>> statements = ['gaussian fit is not the correct color']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([plot_2[0].get_color()])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: gaussian fit is not the correct color
success_message: gaussian fit is the correct color
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_11'
>>> max_score = 3
>>> score = 0
>>> conditions = [plot_1.axes.get_legend() is not None]
>>> statements = ['The legend is not correct on the plot']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str([plot_1.axes.get_legend() is not None])}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: The legend is not correct on the plot
success_message: The legend is correct on the plot
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_12'
>>> max_score = 6
>>> score = 0
>>> conditions = [all(plot_2[0].get_xdata() == x_data), all(plot_2[0].get_ydata() == y_fit)]
>>> statements = ['the x data for plot_2 is not correct', 'the y data for plot_2 is not correct']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: plot_2 does not have the correct values.
success_message: plot_2 is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import submit_score, get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> scorer = submit_score()
>>> question_id = 'q9_13'
>>> max_score = 6
>>> score = 0
>>> conditions = [np.array_equal(plot_1.get_offsets()[:, 1], y_noisy), np.array_equal(plot_1.get_offsets()[:, 0], x_data)]
>>> statements = ['the y data for plot_1 is not correct', 'the x data for plot_1 is not correct']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: plot_1 does not have the correct values.
success_message: plot_1 is correct.
